<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# CP101: Assignment 2 - Control Flow, Functions, and Debugging

**Due Date:** Sep 28, 2026 - 11:59 pm

**Total Points:** 100 (+ 5 Bonus)

---

### Instructions
- Complete all questions in this notebook
- Complete all code cells, then run **Kernel → Restart & Run All** before submitting
- Submit the autograder zip to Gradescope. Every question in this assignment is autograded, so there is no separate written-work PDF.

```python
# Initialize Otter
import otter
grader = otter.Notebook("CP101_A2.ipynb")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Question 1: Code Block Ordering (10 points)

For each question, the code lines are scrambled. If two lines are independent of each other, so that their order does not affect the result, list them in the order they appear in the code block below.

**Task:** Create a variable `q1_1_order` that references letters in the correct order, as a list of strings. A list of strings is written like `["X", "Y", "Z"]`.

### Q1.1: Park Accessibility Report (5 points)

These lines count how many neighborhoods have zero parks. What is the correct order so that the number of neighborhoods with no parks is printed correctly?

```python
Line A:      if parks == 0:
Line B:  parks_by_neighborhood = [3, 0, 1, 5, 2, 1, 0]
Line C:  print(f"Neighborhoods with no parks: {no_park_count}")
Line D:  for parks in parks_by_neighborhood:
Line E:          no_park_count += 1
Line F:  no_park_count = 0
```

```python
q1_1_order = ...
```

```python
grader.check("q1_1")
```

### Q1.2: Transit Frequency Categories (5 points)

These lines count how many bus routes are:
- very frequent: 10 minutes or less
- frequent: more than 10 but 15 minutes or less
- infrequent: more than 15 minutes

Create a variable `q1_2_order` that references a list of letters in the correct order.

```python
Line A:  frequent = 0
Line B:  bus_frequencies = [5, 15, 8, 30, 12, 45]
Line C:      if frequency <= 10:
Line D:  print(f"Very frequent: {very_frequent}, Frequent: {frequent}, Infrequent: {infrequent}")
Line E:          very_frequent += 1
Line F:  for frequency in bus_frequencies:
Line G:  very_frequent = 0
Line H:      elif frequency <= 15:
Line I:          infrequent += 1
Line J:          frequent += 1
Line K:  infrequent = 0
Line L:      else:
```

```python
q1_2_order = ...
```

```python
grader.check("q1_2")
```

---

## Question 2: Bugs + Debugging (22 points)

### Q2.1: Street Tree Size Categories (9 points)

The city classifies street trees by trunk diameter:

| Category | Trunk diameter |
|---|---|
| small | under 12 inches |
| medium | 12 up to but not including 20 inches |
| large | 20 inches or more |

Fix all bugs in the code below so it runs and sorts each tree into the right category.

Expected output: `Small: 3, Medium: 3, Large: 1`

Not every bug here stops the code from running. Some are syntax errors that Python will refuse outright, but one is a logic error in the categories themselves, and that one is the reason the counts come out wrong rather than the reason the code crashes. Check each branch against the table above.

```python
tree_diameters = [8, 15, 22, 6, 18, 19, 10]
small_trees == 0
medium_trees == 0
large_trees == 0
for diameter in tree_diameters:
    if diameter <= 12:
        small_trees += 1
    elif diameter < 20:
        medium_trees = 1
    else
        large_trees += 1
print(f"Small: {small_trees}, Medium: {medium_trees}, Large: {large_trees}")
```

```python
tree_diameters = [8, 15, 22, 6, 18, 19, 10]  # inches
```

```python

```

```python
grader.check("q2_1")
```

### Q2.2: Average Task Time (13 points)

**Part A (4 points):** The following code calculates the total amount of time spent on a series of tasks given a list of `task_times` and then calculates the average amount of time spent on a task. Add `print()` statements for `total` at the three marked locations in the cell below and run it. The prints are not graded; they are how you find the answer.

Watching `total` change on each pass tells you which statement is wrong.

**A1.** Exactly one statement below is the bug. Set `q2_bug` to the **choice number** for that statement.

```text
1. `total = time` replaces the running total each pass instead of adding to it.
2. `count += 1` counts one too many, because it also counts the list itself.
3. `average = total / count` divides the wrong way round.
4. `total = 0` should have started at 1.
5. The loop walks `task_times` itself instead of its positions, so it never sees the values.
```

**A2.** Set `q2_wrong_average` to the average the buggy code actually prints.

```python
# PART A: ADD PRINT STATEMENTS WHERE MARKED (do not change the logic yet)
task_times = [45, 90, 30, 120, 60]  # minutes per task

total = 0
count = 0

# PRINT 1 HERE

for time in task_times:
    total = time
    # PRINT 2 HERE
    count += 1

average = total / count

# PRINT 3 HERE

print(f"Average task time: {average} minutes")
```

```python
q2_bug = ...
q2_wrong_average = ...
```

```python
grader.check("q2_2a")
```

**Part B (9 points):** Write the corrected code using the same `task_times` list.

Expected correct average: 69.0 minutes

```python
task_times = [45, 90, 30, 120, 60]
total = 0
count = 0
for time in task_times:
    total = ...
    count += 1
average = ...
print(...)
```

```python
grader.check("q2_2b")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Question 3: Truthiness Check (10 points)


So far, your conditions have mostly been comparisons, such as:

```python
if population > 300000:
    ...
```

But as discussed, each value in Python has a **truthiness**, meaning that Python can evaluate it as a Boolean (`True` or `False`). Because of this, values can be used directly in a conditional statement.

For the values we have seen so far, Python follows a simple rule:

* `False` and `0` act as **false**.
* Empty sequences or collections, such as `""` and `[]`, act as **false**.
* Everything else acts as **true**.


For each row in the table below:

* If there is anything in the `Given` column, it is placed on the first line.
* The expression in the `condition` column is then placed right after `if`:

```python
# Code from the Given column appears here, if any.

if condition:
    print("This line prints")
```


**Without running the code**, predict whether the `print()` statement runs.

| Expression | Given                       | `condition`    |
| -----------| --------------------------- | -------------- |
| A    | —                           | `0`            |
| B    | —                           | `""`           |
| C    | —                           | `"False"`      |
| D    | —                           | `[]`           |
| E    | —                           | `[0]`          |
| F    | `bike_lanes = 0`            | `bike_lanes`   |
| G    | `neighborhood = "Downtown"` | `neighborhood` |


Now, create a variable named `q3_predictions` that references a **list of seven strings**, one for each part **A through G, in order**.

Each string must be exactly `"prints"` or `"does not print"`.

For example, if Part A prints and Part B does not print, your list would begin:

```python
q3_predictions = ["prints", "does not print", ...]
```

Add one prediction to the list for each expression in the table.

```python
q3_predictions = ...
```

```python
grader.check("q3")
```

For Expression **B** (`if "":`) and **E** (`if [0]:`), pick the reason your prediction was right.

```text
1. `""` is an empty string, and an empty sequence acts as false, so the condition fails.
2. `""` acts as false because the characters spell nothing meaningful.
3. `[0]` acts as true because the list is non-empty; what sits inside it does not matter.
4. `[0]` acts as false because the single value inside it is 0.
5. `[0]` acts as true because 0 is a truthy number.
```

Set `q3_why_b` and `q3_why_e` to the **choice number** for the reason that applies to each.

```python
q3_why_b = ...
q3_why_e = ...
```

```python
grader.check("q3_explain")
```

---

## Question 4: Fill-in-the-Blanks (14 points)

### Q4.1: Bike Lane Quality Assessment (9 points)

The following code creates three buckets for bike lane scores and then takes a list of scores and counts how many fall into each bucket. 

Fill in the blanks so the code categorizes bike lanes correctly.

Expected: `Excellent: 3, Good: 1, Needs improvement: 3`

```python
bike_lane_scores = [85, 45, 92, 38, 67, 88, 52]
```

```python
excellent = _____
good = _____
needs_improvement = _____

for score in bike_lane_scores:
    if score >= 80:
        excellent += 1
    elif score >= 60:
        good += 1
    else:
        needs_improvement += 1

print(f"Excellent: {excellent}, Good: {good}, Needs improvement: {needs_improvement}")
```

```python
grader.check("q4_1")
```

### Q4.2: Finding Your Match (5 points)

Fill in the blanks to search a tuple for a street name.

Then run the **same** search a second time for `absent_street`, which is deliberately not in the tuple, and leave the answer in `found_absent`. A search that only ever reports success is not a search, so both cases have to work.

Expected output:
```
Found it: Mission St
Broadway found? False
```

```python
streets = ("Market St", "Valencia St", "Mission St", "Polk St")
looking_for = "Mission St"
absent_street = "Broadway"
```

```python
found = _____

for _____ in _____:
    if _____ == _____:
        found = _____
        print(f"Found it: {street}")

if not found:
    print(f"{looking_for} was not found")

# The same search again, this time for a street that is not in the tuple.
found_absent = _____

for _____ in _____:
    if _____ == absent_street:
        found_absent = _____

print(f"{absent_street} found? {found_absent}")
```

```python
grader.check("q4_2")
```

---

## Question 5: Write Code from Scratch (27 points)

### Q5.1: Housing Affordability Scanner (9 points)

Write code that:
1. Creates `median_incomes = [45000, 120000, 68000, 95000, 52000]`
2. Counts incomes below $70,000 in `affordable_count`
3. Prints the required message

Expected: `Affordable neighborhoods (income < $70,000): 3`

```python
median_incomes = ...
affordable_count = ...
# your loop here
print(...)
```

```python
grader.check("q5_1")
```

### Q5.2: Password Strength Checker (9 points)

A password is **strong** if it is at least 8 characters long AND contains at least one digit;
otherwise it is **weak**. Use `char.isdigit()` to check each character.

Point each of these four variables at the right value, then print `strength_label` on its
own line:

- `password_length`: how many characters the password has
- `has_digit`: `True` if any character is a digit, otherwise `False`
- `is_strong`: `True` only if the password is both long enough and contains a digit
- `strength_label`: exactly `"Strong password"` or `"Weak password"`

Note the name `has_digit` rather than `is_digit`. The variable does not answer "is this a
digit?", it answers "did a digit turn up anywhere?", and the name should say which. You will notice how much work
this check is to repeat for a second password. Question 6 is the answer to that.

```python
password = "planning2026"
```

```python
password_length = ...
has_digit = ...
# your loop here
is_strong = ...
strength_label = ...
print(strength_label)
```

```python
grader.check("q5_2")
```

### Q5.3: POI Access Score (9 points)

Each amenity type is worth a different number of points toward a walkability score:

| Amenity | Points |
|---|---|
| `park` | 3 |
| `library` | 3 |
| `grocery` | 5 |
| `cafe` | 2 |
| `bus_stop` | 4 |
| anything else | 0 |

Sum the points for every amenity in `amenities` and print the total. Expected: `PoI score: 22 points`

Read the last row of the table before you write the loop. One of the amenities in the list is
not in the table at all, and it has to score 0 rather than crash or be skipped silently. Real
amenity data always contains categories your rules did not anticipate.

```python
amenities = ["cafe", "park", "grocery", "park", "bus_stop", "cafe", "library", "pharmacy"]
```

```python
poi_score = ...
# your loop here
print(...)
```

```python
grader.check("q5_3")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Question 6: Writing Functions (17 points)

### Q6.1: Choosing Function Arguments (8 points)

A city inspector checks whether a park's accessible entrance ramp meets the ADA minimum slope requirement. Right now the check only works for **one** specific ramp:

```python
ramp_run = 12      # feet, measured horizontally
ramp_rise = 1      # feet of vertical climb
max_allowed_ratio = 12   # ADA's 1:12 rule requires run / rise >= 12

if ramp_run / ramp_rise >= max_allowed_ratio:
    print("Ramp meets ADA slope requirements.")
else:
    print("Ramp does NOT meet ADA slope requirements.")
```

ADA's 1:12 rule is about **run** (the horizontal distance the ramp covers) over **rise** (the
height it climbs), not the sloped length of the ramp surface. Naming the parameter `run`
rather than `length` is the difference between a function a colleague can use correctly and
one they will quietly misuse.

**Task:** Turn this into a function `meets_ada_slope(run, rise)` that works for **any** ramp,
and returns `True` or `False` instead of printing. The ADA ratio (`12`) does not change from
ramp to ramp, so it should stay a fixed value inside the function, not a parameter.

```python
def meets_ada_slope(run, rise):
    ...
```

```python
grader.check("q6_1")
```

Which of these explain why `run` and `rise` became **parameters** while the ADA ratio (`12`) did not?

```text
1. `run` and `rise` differ for every ramp, so the caller has to supply them.
2. The ADA ratio is fixed by the regulation, so it belongs inside the function.
3. `12` cannot be a parameter, because Python does not allow numbers as defaults.
4. `run` and `rise` must be parameters because a Python function needs at least two.
5. Making the ratio a parameter would let a caller quietly change the legal standard.
```

Set `q6_reasons` to a **list** of the **choice numbers** for the true statements, in any order.

```python
q6_reasons = ...
```

```python
grader.check("q6_explain")
```

### Q6.2: A Function You Can Reuse (9 points)

Look back at **Q1.2**. You hand-ordered an `if`/`elif`/`else` chain that sorted bus routes into
very frequent, frequent, and infrequent. Then in **Q5.2** you wrote a check you would have to
copy and paste in full to run it on a second password.

That is the problem functions solve. Write the rule **once**, give it a name, and call it as
many times as you like.

**Part 1.** Write a function `categorize_frequency(minutes)` that returns a string:

| Headway | Returns |
|---|---|
| 10 minutes or less | `"very frequent"` |
| more than 10, up to and including 15 | `"frequent"` |
| more than 15 | `"infrequent"` |

**Part 2.** Use your function in a loop over `bus_frequencies` to build a list called
`frequency_labels`, with one label per route, in the same order as the routes.

Pay attention to `10` and `15` themselves. Those two values sit exactly on the boundaries, and
whether you write `<` or `<=` decides which category they land in. Boundary values are where
this kind of rule almost always breaks.

```python
bus_frequencies = [5, 15, 8, 30, 12, 45]  # minutes between buses
```

```python
def categorize_frequency(minutes):
    ...

frequency_labels = ...
# your loop here
print(frequency_labels)
```

```python
grader.check("q6_2")
```

---

## OPTIONAL BONUS (5 points): Code Tracing

Without running the code, trace what it does:

```python
points = 0
actions = [5, -2, 8, -3, 6]

for change in actions:
    if change >= 0:
        points = points + change
    else:
        points = points + (change * 2)
```

Set:
- `bonus_before` to the list of `points` values **before** each of the 5 iterations
- `bonus_after` to the list of `points` values **after** each iteration
- `bonus_final_points` to the final total

```python
bonus_before = ...
bonus_after = ...
bonus_final_points = ...
```

```python
grader.check("bonus_trace")
```

According to the code, how do the points change differently for positive and negative numbers?

```text
1. Positive changes are added as they are, while negative changes are doubled before being added.
2. Positive and negative changes are both doubled.
3. Negative changes are halved before being added.
4. Negative changes are ignored entirely.
5. The doubling makes a loss count twice as heavily as a gain of the same size.
```

Set `bonus_reasons` to a **list** of the **choice numbers** for the true statements, in any order.

```python
bonus_reasons = ...
```

```python
grader.check("bonus_explain")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Make sure you have saved your notebook before running this cell. Then run this cell and click on the download link to get your submission zip.

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for bonus_explain (points: 2)
```
>>> assert not isinstance(bonus_reasons, str), 'Give bonus_reasons as a list of numbers, e.g. [1, 3].'
>>> assert (set(bonus_reasons) if hasattr(bonus_reasons, '__iter__') else {bonus_reasons}) == {1, 5}
```


Tests for bonus_trace (points: 3)
```
>>> assert bonus_before == [0, 5, 1, 9, 3]
```
```
>>> assert bonus_after == [5, 1, 9, 3, 9]
```
```
>>> assert bonus_final_points == 9
```


Tests for q1_1 (points: 5)
```
>>> assert q1_1_order in (['B', 'F', 'D', 'A', 'E', 'C'], ['F', 'B', 'D', 'A', 'E', 'C'])
```


Tests for q1_2 (points: 5)
```
>>> inits, tail = (q1_2_order[:4], q1_2_order[4:])
>>> assert sorted(inits) == ['A', 'B', 'G', 'K'], 'the first four lines are the four initializers'
>>> assert tail == ['F', 'C', 'E', 'H', 'J', 'L', 'I', 'D']
```


Tests for q2_1 (points: 9)
```
>>> assert small_trees == 3
```
```
>>> assert medium_trees == 3
```
```
>>> assert large_trees == 1
```


Tests for q2_2a (points: 4)
```
>>> assert q2_bug == 1
```
```
>>> assert q2_wrong_average == 12.0
```


Tests for q2_2b (points: 9)
```
>>> assert average == 69.0
```
```
>>> assert total == 345
```


Tests for q3 (points: 7)
```
>>> expected = ['does not print', 'does not print', 'prints', 'does not print', 'prints', 'does not print', 'prints']
>>> assert q3_predictions == expected
```


Tests for q3_explain (points: 3)
```
>>> assert q3_why_b == 1
```
```
>>> assert q3_why_e == 3
```


Tests for q4_1 (points: 9)
```
>>> assert excellent == 3
```
```
>>> assert good == 1
```
```
>>> assert needs_improvement == 3
```


Tests for q4_2 (points: 5)
```
>>> assert found is True, 'the Mission St search should have succeeded'
>>> assert found_absent is False, 'searching for a street not in the tuple must leave found_absent False'
```


Tests for q5_1 (points: 9)
```
>>> assert affordable_count == 3
```


Tests for q5_2 (points: 9)
```
>>> assert password_length == 12
>>> assert has_digit is True
>>> assert is_strong is True
>>> assert strength_label == 'Strong password'
>>> assert strength_label == ('Strong password' if is_strong else 'Weak password')
```


Tests for q5_3 (points: 9)
```
>>> assert poi_score == 22
```


Tests for q6_1 (points: 6)
```
>>> assert meets_ada_slope(12, 1) == True
```
```
>>> assert meets_ada_slope(6, 1) == False
```
```
>>> assert meets_ada_slope(24, 2) == True
```
```
>>> assert meets_ada_slope(11, 1) == False
```


Tests for q6_2 (points: 9)
```
>>> assert categorize_frequency(10) == 'very frequent'
>>> assert categorize_frequency(15) == 'frequent'
```
```
>>> assert categorize_frequency(4) == 'very frequent'
>>> assert categorize_frequency(11) == 'frequent'
>>> assert categorize_frequency(16) == 'infrequent'
```
```
>>> assert frequency_labels == ['very frequent', 'frequent', 'very frequent', 'infrequent', 'frequent', 'infrequent']
```


Tests for q6_explain (points: 2)
```
>>> assert not isinstance(q6_reasons, str), 'Give q6_reasons as a list of numbers, e.g. [1, 3].'
>>> assert (set(q6_reasons) if hasattr(q6_reasons, '__iter__') else {q6_reasons}) == {1, 2, 5}
```
