<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# CP101: Assignment 1 - Variables, Data Types & Operators

**Due Date:** Sep 16, 2026 - 11:59 pm

**Total Points:** 100 (+ 2 bonus)

---

### Instructions
- Complete all questions in this notebook
- Complete all code cells, then run **Kernel → Restart & Run All** before submitting
- Submit the autograder zip (coding questions) and the written-work PDF (manual questions) to their respective Gradescope portals

```python
# Initialize Otter
import otter
grader = otter.Notebook("CP101_A1.ipynb")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Question 1: Variable Naming Detective (14 points)
You're reviewing code from a colleague analyzing street widths. Some of their variable names are valid in Python and others are not, several having been pasted straight from spreadsheet column headers.

**Part A (7 points):** For each variable name below, set the corresponding `q1_*` variable to exactly `"VALID"` or `"INVALID"`. Judge each name on its own: the answers are not in any pattern, and valid and invalid names do not alternate.

```
a) street_width_ft
b) StreetWidth
c) 2nd-ave width
d) _private_width
e) street-width
f) 3rd street
g) MAX_LANE_WIDTH
h) class
```

```python
q1_a = ...
q1_b = ...
q1_c = ...
q1_d = ...
q1_e = ...
q1_f = ...
q1_g = ...
q1_h = ...
```

```python
grader.check("q1a")
```

**Part B (7 points):** Spotting that a name is invalid is one skill; explaining *why* is another.

**B1.** Go back through the eight names in Part A and record *why* each invalid name fails, choosing from the statements below. A name can be invalid in more than one way at once, so give **every** reason that applies, not just the first one you spot.

```text
1. It contains a digit somewhere, and Python names may not contain digits at all.
2. It contains `-`, which Python reads as the subtraction operator instead of as part of the name.
3. It uses capital letters, and Python names must be entirely lowercase.
4. It is a reserved Python keyword, so the language has already claimed that word.
5. It starts with a digit, and a Python name must begin with a letter or an underscore.
6. It starts with an underscore, which Python does not allow at the start of a name.
7. It contains a space, and Python reads the space as the end of the name.
8. It is the name of a built-in Python function, so assigning to it is a syntax error.
```

Some of these statements describe actual rules for Python names; others describe common misconceptions. Before using a statement as a reason, decide whether the rule itself is true.

For each letter, set `q1_why_a` through `q1_why_h` to a **list** containing the **choice numbers** for all reasons that apply. The order does not matter: use `[3, 6]` if two reasons apply and `[1]` if only one applies. If you marked a name `"VALID"` in Part A, use the empty list `[]`.

**B2.** `Valid` is not the same as `Pythonic`. Python will accept names that do not follow the community style conventions in PEP 8.

Exactly one of the names you marked `"VALID"` uses `CapWords`, the convention PEP 8 recommends for class names rather than ordinary variable names. Set `q1_not_pythonic` to that name's letter, as a lowercase string such as `"a"`.


```python
q1_why_a = ...
q1_why_b = ...
q1_why_c = ...
q1_why_d = ...
q1_why_e = ...
q1_why_f = ...
q1_why_g = ...
q1_why_h = ...
q1_not_pythonic = ...
```

```python
grader.check("q1b")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Question 2: Put the Script Back in Order (14 points)
A colleague sends you the lines of a short script. It is supposed to work out the **population density** of a neighborhood, in people per square mile, and print it to one decimal place.

The lines arrived scrambled, and two lines that do not belong in the script got mixed in with them.

**The numbers the script works from:**
- Total population: 45,000 people
- Area: 2.3 square miles

```python
1  population_density = total_population / area_square_miles
2  total_population = 45000
3  area_square_miles = 2.3
4  print(f"Population density: {population_density:.1f} people/sq mi")
5  population_density = total_population * area_square_miles
6  print("Population density:", population_density)
```

Work out which four lines belong, then put them in the order Python has to run them. Order is what makes this work: a line that uses a variable only runs correctly if the line that creates that variable has already run.

Record your answer in `q2_order` as a **tuple** holding those four line numbers, first to last. A tuple of numbers is written like `(3, 1, 2)`. In Question 1 you used a list and the order of the numbers did not matter; here the order is the whole answer, and a tuple is the natural way to write a fixed sequence you do not intend to change.

If two lines could correctly go in either order relative to each other, put them in ascending numeric order.

> Referring to code by line number is worth getting used to. It is how Python reports errors back to you, and how a colleague will point you at a problem in a file neither of you is looking at together.

```python
q2_order = ...
```

```python
grader.check("q2_order")
```

**Explain:** You left two lines out of the script. Say which, and what each one would have got wrong.

Set `q2_unused` to a **list** of the two line numbers you did not use.

Then, for each of those two lines, pick the reason it does not belong from this list:

```text
1. It multiplies where it should divide, so the result is not a density at all.
2. It prints the number without rounding it to one decimal place, so it misses the spec.
3. It refers to a variable that no line in the list ever creates.
4. It would stop the script with a ZeroDivisionError.
```

Set `q2_why_5` and `q2_why_6` to the **choice number** for the reason that applies to each line.

```python
q2_unused = ...
q2_why_5 = ...
q2_why_6 = ...
```

```python
grader.check("q2_unused")
```

**Now write the script out for real,** in the order you chose, and run it.

```python
total_population = ...
area_square_miles = ...
population_density = ...
print(...)
```

```python
grader.check("q2_code")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Question 3: Building Permit Calculator (13 points)
Calculate building permit costs using these rules:

**Fees:**
- Base fee: $500
- Area fee: $2.50 per square foot
- Rush processing: an extra $250, charged only when `rush_processing` is `True`

Use the given variables. Compute `total_cost` and print a formatted breakdown. Only `total_cost` is autograded; the exact wording of the print statement is up to you.

You have not met `if` statements yet, so here is the trick for the rush fee. A boolean converts straight to a number: `int(True)` is `1` and `int(False)` is `0`. Multiplying the $250 by that conversion charges the fee when the flag is on and charges nothing when it is off, with no `if` needed.

```python
# Given variables - DO NOT MODIFY
project_area_sqft = 1200
rush_processing = True
```

```python
base_fee = ...
cost_per_sqft = ...
area_cost = ...
rush_fee = ...
total_cost = ...
print(...)
```

```python
grader.check("q3")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


## Question 4: The Type Conversion Problem (14 points)
Complete the code below so that it produces:

```text
Total daily riders: 12750
Average per hour: 531.25
At capacity: False
```




```python
morning_riders = "8500"
afternoon_riders = "4250"
hours_of_operation = 24
max_capacity = 15000
total_riders = ...
average_per_hour = ...
at_capacity = ...
print(...)
print(...)
print(...)
```

```python
grader.check("q4_code")
```

**Explain:** What was wrong, and what fixed it?

**Part 1.** Exactly one statement below correctly describes the problem. Set `q4_problem` to the **choice number** for that statement.

```text
1. The rider counts are strings, so `+` joins them end to end instead of adding them.
2. The rider counts are floats, so the total picks up a decimal point.
3. `hours_of_operation` is stored as text, so the division cannot run.
4. `max_capacity` is not defined, so the comparison has nothing to compare against.
5. The rider counts are negative, so the total comes out too low.
```

**Part 2.** Which conversion produces the expected total as the integer `12750`, with no decimal point? Set `q4_fix` to the **choice number** for the correct conversion.

```text
1. float()
2. int()
3. str()
4. round()
```


```python
q4_problem = ...
q4_fix = ...

```

```python
grader.check("q4_explain")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Question 5: Operator Precedence & Logic (24 points)
A city calculates affordable housing requirements for two proposed buildings:

**Policy:**

* 20% of units must be affordable, for every building regardless of size.
* If a shortfall exists, the developer pays $50,000 per missing unit.
* The city adds a 5% administrative surcharge to the total fee.
* A building that already exceeds its requirement owes nothing. It does not earn a refund.

Find and fix **all** bugs in the code below. Everything you need is arithmetic and comparison operators, parentheses, and `max()`.

Your corrected code should produce:

```text id="0yc3wv"
Building A - In-lieu fee owed: $367,500
Building A - Must pay a fee: True
Building B - In-lieu fee owed: $0
Building B - Must pay a fee: False
```


```python
# BUGGY CODE - Find and fix ALL errors
fee_per_unit = 50000
required_percentage = 0.20

# Building A: a small infill project
total_units_a = 85
affordable_units_included_a = 10
required_affordable_a = total_units_a * required_percentage
shortfall_a = required_affordable_a - affordable_units_included_a
in_lieu_fee_a = shortfall_a * fee_per_unit
needs_to_pay_a = shortfall_a > 0 and total_units_a >= 100

# Building B: a project that already has extra affordable units
total_units_b = 105
affordable_units_included_b = 25
required_affordable_b = total_units_b * required_percentage
shortfall_b = required_affordable_b - affordable_units_included_b
in_lieu_fee_b = shortfall_b * fee_per_unit
needs_to_pay_b = shortfall_b > 0 and total_units_b >= 100

print(f"Building A - In-lieu fee owed: ${in_lieu_fee_a:,.0f}")
print(f"Building A - Must pay a fee: {needs_to_pay_a}")
print(f"Building B - In-lieu fee owed: ${in_lieu_fee_b:,.0f}")
print(f"Building B - Must pay a fee: {needs_to_pay_b}")
```

```python
fee_per_unit = 50000
required_percentage = 0.20
administrative_surcharge = 0.05

total_units_a = 85
affordable_units_included_a = 10
required_affordable_a = ...
shortfall_a = ...
in_lieu_fee_a = ...
needs_to_pay_a = ...

total_units_b = 105
affordable_units_included_b = 25
required_affordable_b = ...
shortfall_b = ...
in_lieu_fee_b = ...
needs_to_pay_b = ...

print(...)
print(...)
print(...)
print(...)
```

```python
grader.check("q5_code")
```

### Part A: Bug Report

Which bugs were actually in the code? Some of the statements below describe problems that were **not** present.

```text
1. `needs_to_pay` included an extra `and total_units >= 100` condition that the policy does not have.
2. `shortfall` was not capped at zero, so a building with more affordable units than required could produce a negative fee.
3. The 5% administrative surcharge was added incorrectly because the expression did not group `1 + administrative_surcharge` together.
4. `required_affordable` used `+` instead of `*`.
5. `fee_per_unit` was set to the wrong dollar amount.
6. The two buildings' values were swapped, so Building A used Building B's numbers.
7. `in_lieu_fee` divided by `fee_per_unit` instead of multiplying by it.
```

Set `q5_bugs` to a **list** of the **choice numbers** for all the real bugs, in any order.


```python
q5_bugs = ...

```

```python
grader.check("q5_a")
```

### Part B: Logical Operator Analysis

Suppose the condition used `or` instead of `and`:

```python
needs_to_pay = shortfall > 0 or total_units >= 100
```

Using the corrected values below, work out what this expression would produce for each building:

* Building A: `shortfall = 7`, `total_units = 85`
* Building B: `shortfall = 0`, `total_units = 105`

Set `q5_or_a` and `q5_or_b` to the `True` or `False` produced for each building.

Then set `q5_or_wrong` to a **list** containing the buildings for which this result would be wrong according to the policy, using `"A"` and `"B"`. If neither is wrong, use the empty list `[]`.



```python
q5_or_a = ...
q5_or_b = ...
q5_or_wrong = ...

```

```python
grader.check("q5_b")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Question 6: Representing Data with Python Types (10 points)

Here is one row of a traffic safety dataset, shown the way it might appear in a report:

**Broadway & 96th Street**

* Crashes last year: 12
* Daily traffic: 18,750 vehicles
* Has traffic signal: Yes
* Safety rating: 2.3 (scale 1–5)

Store the intersection name and these four facts in the variables below.

Four basic Python types appear in this row: text, whole numbers, a decimal value, and a yes/no value. Choose the appropriate Python representation for each.

Two details need special attention:

* Write the traffic count as `18750` (or `18_750`). In Python, `18,750` is not a single integer; the comma acts as a separator.
* Represent the signal value as `True`, not as the string `"Yes"`.

The tests check the **type** as well as the value, so `12.0` will not pass where `12` belongs.


```python
intersection_name = ...
crashes_last_year = ...
daily_traffic_vehicles = ...
has_traffic_signal = ...
safety_rating = ...

```

```python
grader.check("q6")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Question 7: Whole Numbers and Leftovers (11 points)
A surface parking lot has **47 stalls**, and the striping crew lays them out in **rows of 8**.

Plain division is the wrong tool here. `47 / 8` is `5.875`, and there is no such thing as
0.875 of a row. Two operators answer this exactly:

- `//` is **floor division**: it divides and discards the remainder. `17 // 5` is `3`.
- `%` is the **modulo** operator: it gives you only the remainder. `17 % 5` is `2`.

Together they split a whole number into "how many complete groups" and "what is left over",
which is the shape of a lot of planning arithmetic: units per floor, buses per headway,
parcels per block.

Compute:

- `full_rows`: the number of complete rows of 8 the lot holds
- `leftover_stalls`: the stalls remaining after the last complete row

Then print a one-line summary.

```python
total_stalls = 47
stalls_per_row = 8
full_rows = ...
leftover_stalls = ...
print(...)
```

```python
grader.check("q7")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Reflection (BONUS - 2 points)

Answer one:

1. What was the most confusing concept in this assignment? What helped you understand it?
2. What mistake did you make that taught you something important?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Written Work Submission

Below, you will see two cells. Running the first cell will automatically generate a PDF of all questions that need to be manually graded, and running the second cell will automatically generate a zip with your autograded answers. You are responsible for submitting both the coding portion (the zip) and the written portion (the PDF) to their respective Gradescope portals. **Please save before exporting!**

> **Important: You must correctly assign the pages of your PDF after you submit to the correct gradescope assignment. If your pages are not correctly assigned and/or not in the correct PDF format by the deadline, we reserve the right to award no points for your written work.**

If there are issues with automatically generating the PDF in the first cell, you can try downloading the notebook as a PDF by clicking on `File -> Save and Export Notebook As... -> Webpdf`. If that doesn't work either, you can manually take screenshots of your answers to the manually graded questions and submit one single PDF of your screenshots. Either way, **you are responsible for ensuring your submision follows our requirements, we will NOT be granting regrade requests for submissions that don't follow instructions.**

**You must submit the PDF generated via one of these methods, we will not accept screenshots or Word documents.**

```python
from otter.export import export_notebook
from os import path
from IPython.display import display, HTML
name = 'CP101_A1'
export_notebook(f"{name}.ipynb", filtering=True, pagebreaks=True)
if(path.exists(f'{name}.pdf')):
    display(HTML(f"Download your PDF <a href='{name}.pdf' download>here</a>."))
else:
    print("\n Pdf generation failed, please try the other methods described above")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Make sure you have saved your notebook before running this cell. Then run this cell and click on the download link to get your submission zip.

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1a (points: 7)
```
>>> assert q1_a.strip().upper() == 'VALID'
>>> assert q1_b.strip().upper() == 'VALID'
>>> assert q1_c.strip().upper() == 'INVALID'
>>> assert q1_d.strip().upper() == 'VALID'
>>> assert q1_e.strip().upper() == 'INVALID'
>>> assert q1_f.strip().upper() == 'INVALID'
>>> assert q1_g.strip().upper() == 'VALID'
>>> assert q1_h.strip().upper() == 'INVALID'
```


Tests for q1b (points: 7)
```
>>> assert not isinstance(q1_why_a, str), 'Give q1_why_a as a list of numbers, e.g. [2, 5] or [].'
>>> assert (set(q1_why_a) if hasattr(q1_why_a, '__iter__') else {q1_why_a}) == set(), 'q1_why_a should be [] if you marked a VALID.'
>>> assert not isinstance(q1_why_b, str), 'Give q1_why_b as a list of numbers, e.g. [2, 5] or [].'
>>> assert (set(q1_why_b) if hasattr(q1_why_b, '__iter__') else {q1_why_b}) == set(), 'q1_why_b should be [] if you marked b VALID.'
>>> assert not isinstance(q1_why_d, str), 'Give q1_why_d as a list of numbers, e.g. [2, 5] or [].'
>>> assert (set(q1_why_d) if hasattr(q1_why_d, '__iter__') else {q1_why_d}) == set(), 'q1_why_d should be [] if you marked d VALID.'
>>> assert not isinstance(q1_why_g, str), 'Give q1_why_g as a list of numbers, e.g. [2, 5] or [].'
>>> assert (set(q1_why_g) if hasattr(q1_why_g, '__iter__') else {q1_why_g}) == set(), 'q1_why_g should be [] if you marked g VALID.'
```
```
>>> assert not isinstance(q1_why_c, str), 'Give q1_why_c as a list of numbers, e.g. [2, 5] or [].'
>>> assert (set(q1_why_c) if hasattr(q1_why_c, '__iter__') else {q1_why_c}) == {2, 5, 7}
```
```
>>> assert not isinstance(q1_why_e, str), 'Give q1_why_e as a list of numbers, e.g. [2, 5] or [].'
>>> assert (set(q1_why_e) if hasattr(q1_why_e, '__iter__') else {q1_why_e}) == {2}
```
```
>>> assert not isinstance(q1_why_f, str), 'Give q1_why_f as a list of numbers, e.g. [2, 5] or [].'
>>> assert (set(q1_why_f) if hasattr(q1_why_f, '__iter__') else {q1_why_f}) == {5, 7}
```
```
>>> assert not isinstance(q1_why_h, str), 'Give q1_why_h as a list of numbers, e.g. [2, 5] or [].'
>>> assert (set(q1_why_h) if hasattr(q1_why_h, '__iter__') else {q1_why_h}) == {4}
```
```
>>> assert q1_not_pythonic.strip().lower() == 'b'
```


Tests for q2_code (points: 8)
```
>>> assert abs(population_density - 19565.2) < 0.1
```


Tests for q2_order (points: 3)
```
>>> assert not isinstance(q2_order, str), 'Give q2_order as a tuple of line numbers, e.g. (3, 1, 2).'
>>> assert tuple((int(n) for n in q2_order)) == (2, 3, 1, 4)
```


Tests for q2_unused (points: 3)
```
>>> assert sorted(q2_unused) == [5, 6]
```
```
>>> assert q2_why_5 == 1
```
```
>>> assert q2_why_6 == 2
```


Tests for q3 (points: 13)
```
>>> assert total_cost == 3750
```


Tests for q4_code (points: 7)
```
>>> assert total_riders == 12750
```
```
>>> assert average_per_hour == 531.25
```
```
>>> assert at_capacity is False
```


Tests for q4_explain (points: 7)
```
>>> assert q4_problem == 1
```
```
>>> assert q4_fix == 2
```


Tests for q5_a (points: 6)
```
>>> assert not isinstance(q5_bugs, str), 'Give q5_bugs as a list of numbers, e.g. [1, 2, 3].'
>>> assert (set(q5_bugs) if hasattr(q5_bugs, '__iter__') else {q5_bugs}) == {1, 2, 3}
```


Tests for q5_b (points: 6)
```
>>> assert q5_or_a is True
```
```
>>> assert q5_or_b is True
```
```
>>> assert isinstance(q5_or_wrong, list), 'Give q5_or_wrong as a list, e.g. ["A"].'
>>> assert all((isinstance(s, str) for s in q5_or_wrong)), 'Use building names "A" and "B".'
>>> assert {s.strip().upper() for s in q5_or_wrong} == {'B'}
```


Tests for q5_code (points: 12)
```
>>> assert in_lieu_fee_a == 367500
```
```
>>> assert needs_to_pay_a is True
```
```
>>> assert in_lieu_fee_b == 0
```
```
>>> assert needs_to_pay_b is False
```
```
>>> assert required_affordable_a == 17
```
```
>>> assert shortfall_a == 7
```
```
>>> assert required_affordable_b == 21
```
```
>>> assert shortfall_b == 0
```


Tests for q6 (points: 10)
```
>>> assert intersection_name.strip() == 'Broadway & 96th Street'
```
```
>>> assert isinstance(crashes_last_year, int), 'A count of crashes is whole, so use an int, not a float.'
>>> assert crashes_last_year == 12
```
```
>>> assert isinstance(daily_traffic_vehicles, int), 'Write 18,750 as 18750: no comma, no decimal point.'
>>> assert daily_traffic_vehicles == 18750
```
```
>>> assert has_traffic_signal is True
```
```
>>> assert isinstance(safety_rating, float), 'A rating of 2.3 carries a decimal, so it is a float.'
>>> assert safety_rating == 2.3
```


Tests for q7 (points: 11)
```
>>> assert full_rows == 5
```
```
>>> assert leftover_stalls == 7
```
```
>>> assert full_rows * stalls_per_row + leftover_stalls == total_stalls
```
