<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# CP101: Lab 1 - Introducing JupyterHub (DataHub)
## Learning Objectives:
1. Introduce `.py`, `.ipynb` file formats
2. Define and identify typical Python IDEs
3. Navigate the DataHub (JupyterLab) interface
4. Use markdown formatting for documentation
5. Understand the difference between code and text cells
6. Navigate file directories and use terminal commands
<br>

```python
# Initialize Otter
import otter
grader = otter.Notebook("CP101_LAB01.ipynb")
```

---
## 1. What are `.py` and `.ipynb` Files?
Python code lives in one of two common file formats:
- **`.py` files** are plain-text Python scripts: just code, run top-to-bottom as a single program each time, with no embedded output saved in the file.
- **`.ipynb` files** ("IPython Notebook") are structured documents made of separate cells (markdown text cells and code cells) that can be run individually and in any order against a shared, persistent kernel, with each code cell's output saved back into the file.

This class works almost entirely in `.ipynb` notebooks, since mixing explanation, code, and output in one place is useful for data analysis and for showing your work.

<br>

A companion file, `sample_py_file.py`, is included alongside this notebook in your DataHub file browser. Open it and compare it side-by-side with this `.ipynb` notebook: `sample_py_file.py` is nothing but code, top to bottom, with `#` comments as its only narration and no saved output. This notebook, by contrast, mixes markdown text (like this paragraph), code cells, and their results together.

---

## ❗️ Try it Out ❗️
There's a hidden password tucked into a comment near the bottom of `sample_py_file.py`. Open the file in your DataHub file browser, find it, and copy it into the cell below.

```python
password = ...
```

```python
grader.check("q1")
```

---
## 2. Define and identify typical Python IDEs
### What is Python, and Why Use it?
"Programming languages are the tools we use to write instructions for computers to follow." ([CodeAcademy](https://www.codecademy.com/resources/blog/programming-languages)). Python is one of the most widely used programming languages by developers worldwide. In the 2024 Stack Overflow Developer Survey, JavaScript, HTML/CSS, and Python were the three most-used languages overall, with Python close behind at 51% ([Stack Overflow](https://survey.stackoverflow.co/2024/technology#most-popular-technologies-language)); Python was also the single most *desired* language for developers to learn next that year, overtaking JavaScript in that specific category ([Stack Overflow Blog](https://stackoverflow.blog/2025/01/01/developers-want-more-more-more-the-2024-results-from-stack-overflow-s-annual-developer-survey/)).

<br>

### What is an IDE?
"An integrated development environment (IDE) is a software application that helps programmers develop software code efficiently. It increases developer productivity by combining capabilities such as software editing, building, testing, and packaging in an easy-to-use application. Just as writers use text editors and accountants use spreadsheets, software developers use IDEs to make their job easier" ([Amazon](https://aws.amazon.com/what-is/ide/)).

<br>

### Navigating DataHub and Jupyter Notebooks

There are a wide array of IDE options for coding including PyCharm, Visual Studio Code, and Atom, but for this class we will be working in **DataHub**, UC Berkeley's hosted JupyterHub service. DataHub gives every student the same preconfigured environment (Python, JupyterLab, and every package this course needs) with no local installation, so nobody's progress depends on their own laptop's setup.
<br>

### Learn to Navigate DataHub
By now, you should have completed the following steps:
1. Navigated to [datahub.berkeley.edu](https://datahub.berkeley.edu) (or the course's DataHub link on bCourses)
2. Signed in with your Berkeley CalNet credentials
3. Launched the JupyterLab interface

<br>

---

<br>



## 3. Create Your First Notebook
Now that you have completed these initial steps, I would like you to complete the following steps to learn to navigate the JupyterLab interface, understand how to use markdown, understand the difference between code and text cells, and navigate file directories and use terminal commands.
<br>

1. In the Launcher, under the **Notebook** section, click the **Python 3** tile to create a new notebook. (If the Launcher isn't showing, click the blue **+** button in the top-left toolbar.)
<br>

### Change the name of your notebook
2. Rename your notebook to `"{firstName}_{sid}"`, where `{sid}` is your Student ID. Right-click the file's name in the left-hand file browser and choose **Rename**, or use the **File** menu's rename option.
<br>

**Example:** "`li_3031234567`"

If your first name has multiple parts, join them with underscores too (e.g. "`mary_jane_3031234567`").

<br>

**A note on saving:** JupyterLab autosaves periodically, but use `Ctrl+S` (`Cmd+S` on Mac) to save on demand. Your files persist in your DataHub home directory between sessions, so there's no separate "Drive" to sync to.

<br>

---

## 4. Working with Text and Markdown Cells
### What is a Text Cell?
A text cell contains Markdown-formatted documentation that renders as formatted text when executed, used for explanations and instructions.

### Understanding Markdown Used in Text Cells
**Markdown** is a lightweight markup language that allows you to format text using simple syntax. It's widely used for:
* Documentation
* README files
* Jupyter notebook text cells
* Websites and blogs

<br>

#### Some helpful examples of markdown syntax:
Note that this syntax might vary slightly from platform to platform

<br>

##### Formatting Headers

```markdown
    # Header 1 (Largest)
    ## Header 2
    ### Header 3
    #### Header 4
    ##### Header 5
    ###### Header 6 (Smallest)
```

##### Formatting Text

```markdown
    **Bold text**
    *Italic text*
    ***Bold and italic***
    ~~Strikethrough~~
```

##### Formatting Lists

```markdown
    - Unordered list item 1
    - Unordered list item 2
    - Nested item

    1. Ordered list item 1
    2. Ordered list item 2
      1. Nested ordered item
```

##### Links and Images

```markdown
    [Link text](URL)
    ![Image alt text](image_URL)
```

##### Code

````markdown
`Inline code`

```python
# Code block
print("Hello, World!")

```
````

##### Tables

```markdown
| Column 1 | Column 2 | Column 3 |
|----------|----------|----------|
| Row 1    | Data     | More data|
| Row 2    | Data     | More data|
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## ❗️ Try it Out ❗️
This is your very first "Try it Out" activity!  These activities are both designed to test your knowledge and evaluate the degree to which you understand Lab concepts. In this Try it Out exercise you will be creating a Text Cell.

Add a text cell in your new notebook by following the instructions below.
1. Insert a new cell (press `B` in command mode, or click the **+** button in the toolbar), then set its type to **Markdown** using the cell-type dropdown in the toolbar (or press `M` in command mode)
2. In this cell, write one sentence reflecting on our discussion from earlier comparing `.py` files and `.ipynb` files. In this sentence you might discuss their key differences in structure, execution, and use cases.  

**Example:**

```
`.py` files are plain-text Python scripts: just code (plus `#` comments if you add them), no embedded output, run top-to-bottom as a single program each time. `.ipynb` files are structured documents made of separate cells (markdown text cells and code cells) that you can run individually and in any order against a shared, persistent kernel -- and each code cell's output gets saved back into the file alongside the code that produced it.
```


<br>

---

<br>

## 5. Working with Code Cells
### Understanding Code Cells

- Interactive execution: Run individual cells independently
- Output display: Results appear directly below the cell
- State persistence: Variables remain in memory between cell executions
- Execution order: Cells can be run in any order (though this can cause confusion)

#### More on Execution Order
Jupyter notebooks maintain a kernel state that persists across cell executions.

**This means that...**
- Variables defined in one cell are available in other cells
- Cells can be executed in any order
- The number in brackets [1] shows execution order


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## ❗️ Try it Out ❗️
Add your first code cell by following the instructions below.

1. Enter the following text into a "Code Cell":

  ```
    print('Hello World')
  ```
  *Hint: A new cell can be created with the **+** button in the notebook toolbar, or by pressing `B` (add cell below) while a cell is selected in command mode; its type (code vs. markdown) is set with the cell-type dropdown in the toolbar, or the `Y` / `M` shortcuts. There should be a code cell already present in your existing notebook.*

<br>

2. From the menu bar, click **Run → Run All Cells** to run every cell in the notebook. You can also run just the selected cell with the ▶ button in the toolbar, or `Shift+Enter`.


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 6. Exploring the JupyterLab Interface

### Touring the Interface Elements

#### Menu Bar
| Menu | Functions |
|------|-----------|
| **File** | Create, open, save, rename, download notebooks |
| **Edit** | Cut, copy, paste, find and replace |
| **View** | Toggle panels (file browser, table of contents, etc.) |
| **Run** | Run one cell, run all cells, run all above/below |
| **Kernel** | Restart, interrupt, or change the Python kernel |
| **Tabs** | Switch between open files and panels |
| **Settings** | Theme, keyboard shortcuts, other preferences |

#### Toolbar
| Button | Function |
|--------|----------|
| **Save icon** | Save the notebook |
| **+** | Insert a new cell below |
| **Scissors icon** | Cut the selected cell |
| **▶** | Run the current cell |
| **■** | Interrupt the kernel |
| **↻** | Restart the kernel |
| **Cell-type dropdown** | Switch a cell between Code / Markdown / Raw |

#### Cell Operations & Shortcuts
*(These apply in "command mode" — press `Esc` first, or click the cell's left margin, so the cell border is highlighted rather than the cursor blinking inside it.)*

| Action | Shortcut | Description |
|--------|----------|-------------|
| **Run cell** | `Shift + Enter` | Run and move to next cell |
| **Run cell in place** | `Ctrl + Enter` | Run but stay in current cell |
| **Add cell above** | `A` | Insert new cell above current |
| **Add cell below** | `B` | Insert new cell below current |
| **Delete cell** | `D`, `D` (press twice) | Delete current cell |
| **Change to code** | `Y` | Convert cell to code cell |
| **Change to markdown** | `M` | Convert cell to text/markdown cell |
| **Enter edit mode** | `Enter` | Start typing inside the cell |

#### Left Sidebar Panels
| Panel | Purpose |
|-------|---------|
| **File Browser** | Browse and manage your DataHub files |
| **Running Terminals and Kernels** | See and shut down active sessions |
| **Table of Contents** | Navigate through markdown headers |
| **Extension Manager** | Browse/install JupyterLab extensions (if enabled) |

#### Cell Types
| Type | Description |
|------|-------------|
| **Code cells** | Execute Python code |
| **Markdown cells** | Formatted text/documentation |
| **Raw cells** | Passed through unprocessed (rarely used in this course) |


<br>

---

<br>

## ❗️ Try it Out ❗️
Complete the following tasks:
- Add a new code cell above this text cell
- Change a code cell to a text cell and back
- Use the table of contents to navigate to different sections
- Open the file browser and look at your DataHub home directory
- Try running a cell with `Ctrl+Enter` vs `Shift+Enter`
- Delete a cell and then undo the action (`Ctrl+Z`) to revert your notebook back to its state before this Try it Out exercise.


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 7. Introducing the File Directory and Terminal

#### What is a file directory?
A file directory (or folder) is a way of organizing files in a hierarchical structure on a computer system. Think of it like a filing cabinet:

- Root directory: The main folder (like the filing cabinet itself)
- Subdirectories: Folders within folders (like drawers and folders within drawers)
- Files: Individual documents stored in directories

#### DataHub's File System
On DataHub, you have access to:
  1. Your Home Directory: your personal, persistent storage — everything you save here stays put between sessions, just like a small personal server
  2. Course Materials: any files provided by the course (e.g. a cloned repo or shared folder)
  3. Uploaded Files: files you upload during the session via the file browser

  > *Try opening up your file directory folder.  What do you see?*

#### What is a Terminal?
A terminal (also called command line, command prompt, or shell) is a text-based interface for interacting with your computer's operating system. Instead of clicking icons and menus, you type commands.

##### Common Terminal Commands
This lab only needs the four commands below -- you will use every one of them in the Try It Out that follows. (There are many more shell commands out there, but no need to front-load them all before you have used even these four.)

- `pwd`: Print working directory (show current location)
- **`ls`: List files and directories**
- **`cd`: Change directory**
- **`mkdir`: Create new directory**

<br>

**To open a terminal on DataHub: from the Launcher, click "Terminal" (or use File → New → Terminal from the menu bar).**


<br>

---

<br>

## ❗️ Try it Out ❗️
Complete the following tasks **in your own Notebook.** Copy the output into your Text cell.

1. List Current Directory Contents
Run the `ls` command in the terminal and copy the output into your Text Cell under the header "# LS Output"

2.  Create a new directory by
running the `mkdir` command. The syntax should read as follows:

    ```markdown
    mkdir my_lab_folder
    ```

3. Change your directory by running the `cd` command.  List the contents of the folder by running the `ls` command.  Use this syntax.  Paste the output into your markdown under the header "# New Directory"

```markdown
cd my_lab_folder && ls
```

> **Understanding the Commands**

- `ls`: Lists files and folders in the current directory
- `mkdir my_lab_folder`: Creates a new directory called "my_lab_folder"
- `cd my_lab_folder && ls`: Changes to the new directory and lists its contents (should be empty)


## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload your submission zip to **LAB01 - Introducing JupyterHub (DataHub)** on [Gradescope](https://www.gradescope.com/).

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1 (points: 0)
```
>>> assert password == 'gob ears'
```
