<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# CP101: Lab 3 - Python Basics I

This lab gives you hands-on practice with the building blocks of Python: arithmetic,
values and types, strings, variables, and the small handful of statements you will reach
for in almost every notebook. We will use small pieces of city data (ridership counts,
station names, coordinates) as our running examples.

**How to work:** run each code cell with `Shift + Enter`, change the values to see what
happens, and complete every **Exercise** cell. Most exercises are checked automatically
once you assign your answer to the named variable; a few ask you to choose your own
example (e.g. "a station of your choice") and are checked for type/shape rather than an
exact value.

*The topic sequence in this lab is inspired by Chapters 1-2 of Think Python (Allen B.
Downey); the explanations and examples are written for CP101.*

```python
# Initialize Otter
import otter
grader = otter.Notebook("CP101_LAB03.ipynb")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 1. Python as a calculator

Python evaluates arithmetic with the operators you already know, plus a couple that are
especially handy for counting things.

| Operator | Meaning | Example |
|---|---|---|
| `+` `-` `*` | add, subtract, multiply | `4500 + 380` |
| `/` | division (always a decimal) | `9 / 2` -> `4.5` |
| `//` | floor division (drop the remainder) | `9 // 2` -> `4` |
| `%` | modulo (the remainder) | `9 % 2` -> `1` |
| `**` | power | `2 ** 10` |

```python
# weekday entries at two BART stations
downtown_berkeley = 4521
ashby = 2980

downtown_berkeley + ashby        # total riders across the two stations
```

```python
# a decimal result, even when it divides evenly
84 / 2
```

### Order of operations

Python follows the usual math precedence: powers first, then multiplication and division,
then addition and subtraction. Use parentheses when you want a different order.

```python
print(2 + 3 * 4)      # multiplication happens first -> 14
print((2 + 3) * 4)    # parentheses change the order -> 20
```

### Floor division and modulo

These two are a pair, and they show up constantly in planning problems. `//` tells you
*how many whole groups* you can make, and `%` tells you *what is left over*.

```python
waiting = 1547         # passengers on the platform
capacity = 200         # passengers per train

full_trains = waiting // capacity   # whole trains we can fill
left_over   = waiting % capacity     # passengers waiting for the next one
print("full trains:", full_trains)
print("left over:  ", left_over)
```

---

### Exercise 1

A bus route runs every day of a 4-week period (28 days). If it makes 46 trips per day,
how many trips is that in total? Assign the total to `total_trips`. Then use `%` to find
how many trips are left over if those trips are grouped into shifts of 8, and assign that
to `shift_leftover`.

```python
total_trips = ...
shift_leftover = ...
```

```python
grader.check("q1")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 2. Built-in functions

A **function** takes some input in parentheses (its **arguments**) and gives back a
result. Python has many built in. A few numeric ones:

```python
print(abs(-45))          # distance from zero
print(round(4.5167, 2))  # round to 2 decimal places
print(min(4521, 2980, 6110))   # smallest
print(max(4521, 2980, 6110))   # largest
```

---

### Exercise 2

Three stations recorded `5120`, `4875`, and `5099` entries. Assign the busiest count to `busiest` and the quietest count to `quietest` using `max` and `min`.

```python
busiest = ...
quietest = ...
```

```python
grader.check("q2")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 3. Values and their types

Every value in Python has a **type**. The four you will meet first:

- `int` : a whole number, like `4521`
- `float` : a number with a decimal point, like `37.8715`
- `str` : text, in quotes, like `"Ashby"`
- `bool` : a truth value, either `True` or `False`

Ask Python for a type with the `type()` function.

```python
print(type(4521))        # int
print(type(37.8715))     # float
print(type("Ashby"))     # str
print(type(True))        # bool
```

### Converting between types

Sometimes a value has the wrong type for what you want to do, for example a number that
arrived as text. Convert it with `int()`, `float()`, or `str()`.

```python
entries_text = "4521"          # this is a string, not a number
entries = int(entries_text)     # now it is an int
print(entries * 2)              # 9042, real arithmetic

print("The station had " + str(entries) + " entries.")  # number back to text
```

---

### Exercise 3

The value `"37.8715"` is a latitude stored as text. Convert it to a `float` and assign it to `latitude`. (For comparison, try converting it with `int()` in a scratch cell instead -- notice what you lose, and why that would be a problem for a coordinate.)

```python
latitude = ...
```

```python
grader.check("q3")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 4. Strings

A **string** is a piece of text in single or double quotes. `+` joins strings together
(concatenation) and `*` repeats them.

```python
line = "Richmond"
color = "Red"
print(color + " Line to " + line)     # join with spaces of our own
print("-" * 30)                        # a divider, 30 dashes
```

### Indexing and slicing

The characters of a string are numbered from `0`. Use square brackets to pull out one
character, or a **slice** `[start:stop]` to pull out a range (the `stop` index is not
included). Negative numbers count from the end.

```python
station = "MacArthur"
print(station[0])      # 'M', the first character
print(station[-1])     # 'r', the last character
print(station[0:3])    # 'Mac', characters 0, 1, 2
print(len(station))    # 9, how many characters
```

### String methods

A **method** is a function attached to a value, called with a dot. Strings have many for
cleaning and checking text.

```python
raw = "  downtown BERKELEY  "
print(raw.strip())              # remove the spaces at the ends
print(raw.strip().title())      # 'Downtown Berkeley'
print(raw.upper())              # ALL CAPS
print("Berkeley" in raw)        # True: is that substring present?
print(raw.replace("BERKELEY", "Oakland"))
```

### f-strings: putting values inside text

An **f-string** (note the `f` before the quote) lets you drop variables straight into a
string with `{ }`. It is the cleanest way to build a readable line of output.

```python
name = "Ashby"
entries = 2980
print(f"{name} had {entries} entries today.")
print(f"That is about {entries / 7:.0f} per day on average.")   # :.0f = no decimals
```

---

### Exercise 4

Start from `full_name = "rockridge station"`.
1. Assign its title-cased version (`Rockridge Station`) to `title_cased`.
2. Assign just the first word (`Rockridge`, using a slice -- hint: it is 9 characters long) to `first_word`.
3. Use an f-string to assign `summary_line` so it reads: `Rockridge Station has 9 letters in its first word.`

```python
full_name = "rockridge station"
title_cased = ...
first_word = ...
summary_line = ...
```

```python
grader.check("q4")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 5. Variables

A **variable** is a name that refers to a value. The `=` sign is the **assignment
operator**: it takes the value on the right and binds the name on the left to it. It does
*not* mean "equals" in the math sense.

```python
station_name = "Downtown Berkeley"
daily_entries = 4521
latitude = 37.8698

# variables can be used in expressions, just like the values they hold
annual_estimate = daily_entries * 260     # roughly 260 weekday services a year
print(station_name, "->", annual_estimate)
```

### Naming variables well

Good names make code readable. The rules and the house style:

- start with a letter or `_`, then letters, digits, or `_` (no spaces, no dashes);
- cannot be a Python keyword (`for`, `if`, `class`, ...);
- names are case-sensitive (`Ashby` and `ashby` differ);
- **use `snake_case`**: lowercase words joined by underscores, like `daily_entries`.

A name should say what the value *means*, so `daily_entries` beats `x` or `de`.

```python
# clear, snake_case names
median_income = 96400
has_elevator = True
print(median_income, has_elevator)
```

### Reassignment

A variable can be pointed at a new value at any time. The old value is simply forgotten.

```python
riders = 4521
print(riders)
riders = riders + 100     # update: take the current value, add 100, rebind the name
print(riders)
```

---

### Exercise 5

Create three well-named `snake_case` variables describing a station of your choice: `station_name_choice` (a `str`), `daily_entries_choice` (an `int`), and `has_parking_choice` (a `bool`). Then print each one together with its `type()`. (Since this is your own example, we only check the *types*, not exact values.)

```python
station_name_choice = ...
daily_entries_choice = ...
has_parking_choice = ...
```

```python
grader.check("q5")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 6. Assigning several variables at once

Python lets you assign several names in a single line. This is tidy for related values,
and it gives you a clean way to **swap** two variables.

```python
# one line, three names
north, south, transfer = "Richmond", "Fremont", "MacArthur"
print(north, south, transfer)

# swap two values without a temporary variable
a, b = 10, 20
a, b = b, a
print(a, b)      # 20 10
```

---

### Exercise 6

In one line, assign `lat` and `lon` the values `37.8698` and `-122.2679`. Then, in one more line, swap them (just to see it work) and print both.

```python
lat, lon = ..., ...
lat, lon = lon, lat
```

```python
grader.check("q6")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 7. The `import` statement

Python's core is small on purpose. Extra tools live in **modules** that you bring in with
`import`. The `math` module, for instance, adds mathematical functions and constants.

```python
import math

print(math.sqrt(2))     # square root
print(math.pi)          # the constant pi

# you can import specific names directly, or give a module a shorter alias
from math import floor
print(floor(4.9))       # 4
```

---

### Exercise 7

The area of a circular service zone with radius `r` is `math.pi * r ** 2`. Compute the area of a zone with radius `800` meters, rounded to the nearest whole number, and assign it to `service_area`.

```python
service_area = ...
```

```python
grader.check("q7")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 8. Printing, and expressions vs statements

`print()` displays its arguments. Give it several, separated by commas, and it puts a
space between them. Two useful keyword arguments: `sep` (what to place between items) and
`end` (what to put at the very end, a newline by default).

```python
print("Ashby", "Rockridge", "MacArthur")               # spaces between
print("Ashby", "Rockridge", "MacArthur", sep=" -> ")   # custom separator
print("loading", end="... ")
print("done")                                          # stays on the same line
```

An **expression** is anything that produces a value (`4521 + 100`, `station.upper()`). A
**statement** is an instruction that performs an action (an assignment, an `import`, a
`print`). A cell can hold several statements; only the value of the last bare expression
is shown automatically, which is why we use `print()` to see more than one result.

```python
x = 5 * 5      # a statement: it acts (binds x), it does not display anything
x              # an expression on its own: its value, 25, is displayed
```

---

### Exercise 8

Using `" | ".join(...)`, build the exact line `Red | Orange | Yellow` and assign it to `color_line`, then print it.

```python
color_line = ...
print(color_line)
```

```python
grader.check("q8")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 9. Comments

Anything after a `#` on a line is a **comment**: Python ignores it, but people reading
your code do not. Good comments explain *why* something is done, not just restate the
code.

```python
# convert the platform crowd into whole trainloads
waiting = 1547
capacity = 200
trains = waiting // capacity     # // because a partial train still needs a whole train slot
print(trains)
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Putting it together

---

### Exercise 9 (wrap-up)

Using what you practiced above, for a station of your choice:
1. store its name, daily entries, and latitude in three well-named variables (`station_choice`, `entries_choice`, `latitude_choice`);
2. compute a rough annual entries estimate (entries times `260`) and assign it to `annual_estimate`;
3. print one f-string line that reports the name, the annual estimate, and the latitude rounded to 3 decimal places.

For example your output might read: `Ashby: about 774800 entries a year, at 37.853 N.` (Since this is your own example, we check the *relationship* between your values, not exact numbers.)

```python
station_choice = ...
entries_choice = ...
latitude_choice = ...
annual_estimate = ...
```

```python
grader.check("q9")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Glossary

- **value** : a basic piece of data, such as `4521` or `"Ashby"`.
- **type** : the kind of a value (`int`, `float`, `str`, `bool`).
- **expression** : a combination of values and operators that produces a value.
- **statement** : an instruction that performs an action (assignment, `import`, `print`).
- **variable** : a name that refers to a value.
- **assignment** : binding a name to a value with `=`.
- **function** : a named operation you call with arguments in parentheses.
- **method** : a function attached to a value, called with a dot (`text.upper()`).
- **module** : a file of extra tools you bring in with `import`.
- **comment** : text after `#` that Python ignores.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Acknowledgments

The topic sequence in this lab is **inspired by** Chapters 1-2 of
[*Think Python*, 3rd ed.](https://allendowney.github.io/ThinkPython/) by Allen B. Downey;
the explanations and examples here were written for CP101. 

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload your submission zip to **LAB03 - Python Basics I** on [Gradescope](https://www.gradescope.com/).

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1 (points: 1)
```
>>> assert total_trips == 1288
>>> assert shift_leftover == total_trips % 8
```


Tests for q2 (points: 1)
```
>>> assert busiest == 5120
>>> assert quietest == 4875
```


Tests for q3 (points: 1)
```
>>> assert isinstance(latitude, float)
>>> assert latitude == 37.8715
```


Tests for q4 (points: 1)
```
>>> assert title_cased == 'Rockridge Station'
>>> assert first_word == 'Rockridge'
>>> assert summary_line == 'Rockridge Station has 9 letters in its first word.'
```


Tests for q5 (points: 1)
```
>>> assert isinstance(station_name_choice, str)
>>> assert isinstance(daily_entries_choice, int) and (not isinstance(daily_entries_choice, bool))
>>> assert isinstance(has_parking_choice, bool)
```


Tests for q6 (points: 1)
```
>>> assert lat == -122.2679
>>> assert lon == 37.8698
```


Tests for q7 (points: 1)
```
>>> assert service_area == round(math.pi * 800 ** 2)
```


Tests for q8 (points: 1)
```
>>> assert color_line == 'Red | Orange | Yellow'
```


Tests for q9 (points: 2)
```
>>> assert isinstance(station_choice, str)
>>> assert isinstance(entries_choice, int) and (not isinstance(entries_choice, bool))
>>> assert isinstance(latitude_choice, (int, float))
>>> assert annual_estimate == entries_choice * 260
```
