<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# CP101: Assignment 0 - Gradescope Submission Practice

**Due Date:** Week 1

**Total Points:** 0

---

### Instructions
- Complete both questions in this notebook
- Run all cells before submitting to make sure your code works
- Save your work, then export a PDF (see the submission section at the bottom)
- Download the autograder zip and submit **both** the zip and the PDF on [Gradescope](https://www.gradescope.com/)

This assignment is practice so you can get comfortable with the submission workflow before graded assignments are released.

```python
# Initialize Otter
import otter
grader = otter.Notebook("CP101_A0.ipynb")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Question 1: Twenty Twenty-Six (0 points)

Assign `twenty_twenty_six` to an arithmetic expression that equals `2026`.

Do not use the number `2026` by itself.

```python
twenty_twenty_six = ...
```

```python
grader.check("q1")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Question 2: Gradescope Submission Checklist (0 points)

List the two files you will upload to [Gradescope](https://www.gradescope.com/) and which assignment each goes to.

**_Hint:_** You can find this in the submission instructions below.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Written Work Submission

Below, you will see two cells. Running the first cell will automatically generate a PDF of all questions that need to be manually graded, and running the second cell will automatically generate a zip with your autograded answers. You are responsible for submitting both the coding portion (the zip) and the written portion (the PDF) to their respective Gradescope portals. **Please save before exporting!**

> **Important: You must correctly assign the pages of your PDF after you submit to the correct gradescope assignment. If your pages are not correctly assigned and/or not in the correct PDF format by the deadline, we reserve the right to award no points for your written work.**

If there are issues with automatically generating the PDF in the first cell, you can try downloading the notebook as a PDF by clicking on `File -> Save and Export Notebook As... -> Webpdf`. If that doesn't work either, you can manually take screenshots of your answers to the manually graded questions and submit one single PDF of your screenshots. Either way, **you are responsible for ensuring your submision follows our requirements, we will NOT be granting regrade requests for submissions that don't follow instructions.**

**You must submit the PDF generated via one of these methods, we will not accept screenshots or Word documents.**

```python
from otter.export import export_notebook
from os import path
from IPython.display import display, HTML
name = 'a0'
export_notebook(f"{name}.ipynb", filtering=True, pagebreaks=True)
if(path.exists(f'{name}.pdf')):
    display(HTML(f"Download your PDF <a href='{name}.pdf' download>here</a>."))
else:
    print("\n Pdf generation failed, please try the other methods described above")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload both files on [Gradescope](https://www.gradescope.com/):

1. **A0 - Gradescope Submission Practice CODE:** upload the `.zip` from the export cell below
2. **A0 - Gradescope Submission Practice WRITTEN:** upload the PDF from the cell above

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1 (points: 0)
```
>>> assert twenty_twenty_six == 2026
```
