<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Pandas Fundamentals

### CP101: Introduction to Urban Data Analytics

**Week 5, Wednesday** &nbsp;·&nbsp; Dictionaries and Pandas I: Wrangling

**Estimated time:** 45 minutes &nbsp;·&nbsp; **Not graded, not submitted.**

<hr style="border: 1px solid #fdb515;" />

### How to work

- Run each code cell with `Shift + Enter`, in order, from the top.
- Nothing in this notebook is submitted or graded. The checks are for you.

```python
# Initialize Otter
import otter
grader = otter.Notebook("CP101_W05_S2_pandas_fundamentals.ipynb")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Learning Objectives

By the end of this session, you will be able to:

1. Create and access Pandas `Series` using different construction methods and indexing (`.loc`, `.iloc`)
2. Create DataFrames from dictionaries, CSV files, and API responses
3. Inspect a DataFrame using `.shape`, `.head()`, `.sample()`, `.dtypes`, `.info()`, `.describe()`, and `.columns`
4. Apply vectorized operations, including `.str` accessor methods, instead of writing manual loops

---

##### [Pandas Data Wrangling Cheat Sheet](https://pandas.pydata.org/Pandas_Cheat_Sheet.pdf)

#### Jupyter Keyboard Shortcuts

| Shortcut | Action |
|----------|--------|
| `Shift + Enter` | Run the current cell |
| `Esc + A/B` | Insert a new cell (Above/Below) |
| `Esc + M` | Convert cell to Markdown |
| `Ctrl + S` (`Cmd + S` on a Mac) | Save notebook |

**Data for this notebook** lives in `data/raw/` inside this folder. It is already there, and nothing needs downloading.

`data/raw/` is a read-only reference copy: do not edit or overwrite files there. To modify the data, load it into a variable and work with that copy instead.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 1. What is Pandas?

[Pandas](http://pandas.pydata.org/) is a versatile and popular Python library for data analysis.

The name comes from **pan**el **da**ta, the econometrics term for observations of many units over many time periods, which is the kind of table the library was first written to handle.

### What is NumPy?
`NumPy` is the Python library for numerical arrays. Its `ndarray` stores values of one type in a contiguous block of memory, which is why arithmetic on a whole array runs far faster than a `for` loop over a list. Pandas is built on top of NumPy and adds labels, mixed column types and file readers, for users whose data is not only numbers.


### Easy-to-use data structures
In pandas, data is typically stored in a `DataFrame`: a table with rows and columns, plus row labels and column names, in which each column can hold a different data type. It is the closest thing in Python to a spreadsheet, with one difference that matters: every column has exactly one type.

### Supports data read/write from multiple formats
One of the most useful features of pandas is its ability to read data
from numerous different data formats directly. For example, pandas
supports reading and writing data from/to:

-   CSV
-   JSON
-   HTML
-   MS Excel
-   Stata
-   SAS
-   parquet
-   SQL (Postgresql, MySQL, Oracle, MariaDB, etc.)

You can view the full list of supported data formats from the [pandas
docs](https://pandas.pydata.org/docs/user_guide/io.html).

### Pandas has two main data structures: `Series` and `DataFrame`

#### Import pandas first

```python
# pandas is imported under the short name pd by convention, so every pandas
# function in a notebook is recognisable at a glance as pd.something
import pandas as pd

print(f'Pandas Version: {pd.__version__}')
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 2. Series

A `Series` is a one-dimensional labeled array holding values of one type (integers, strings, floating point numbers, or Python objects). It is the equivalent of a single spreadsheet column.

The axis labels are collectively referred to as the index.

**The basic method to create a `Series` is to call:**

```python
my_series = pd.Series(
    data,
    index=index,
    name=name,
)
```

Here, `data` can take various forms:  

- A Python dictionary  
- An array  
- A scalar value (e.g., `5`)  

The provided **index** is a **list** of axis labels, which can be strings, time objects, or non-sequential numbers.  

A `Series` can also include an optional `name` argument, which serves as the column heading when converted into a DataFrame.

### 2.1 Creating Series

### From `list` - no index passed
If no index is passed, pandas assigns the positions `0, 1, 2, ...` as labels. They appear down the left when the Series is displayed.

```python
series_1 = pd.Series(
    [100, 200, 500, 800, 1200],
    name = "amount" # name of the column
    )

```

To display an object in a notebook, write its name as the last line of a cell.

```python
series_1
```

### From `list` - index passed

```python
series_2 = pd.Series(
    [100, 200, 500, 800, 1200],
    index=['a', 'b', 'c', 'd', 'e'],
    name = "amount" # name of the column
    )
```

```python
# Note the labels on the left and the data type at the bottom
series_2
```

### From `dict`
The keys become the index labels.

```python
data_dict = {'a': 100, 'b': 200,'c': 500, 'd': 800,'e': 1200}
series_3 = pd.Series(
    data_dict,
    name = "amount" # name of the column
    )
series_3
```

### 2.2 Accessing values in Series

### `Series` is dict-like - Accessing through index name

A Series is also like a fixed-size dict in that you can get and set values by index label:

```python
# the value at label 'a'
series_2['a']
```

```python
print(series_3['c'])
print(series_3['d'])
```

### Containment check
The `in` keyword tests membership. As a dictionary checks its keys, a `Series` checks its **index labels**, not its values:

```python
'b' in series_1    # False: the labels of series_1 are 0 to 4, and `in` never looks at the values
```

```python
'b' in series_2    # True: 'b' is one of the labels of series_2
```

### `.loc` and `.iloc`: label-based and position-based indexing

Just as you reach an item in a list by its position, you can reach a value in a `Series` by its **label** or by its **position**. Pandas gives you one tool for each, and the names say which is which:

- `.loc` is short for **loc**ation, and the location it asks for is the index **label**.
- `.iloc` is **i**nteger **loc**ation: the same idea, counting positions from `0`.

Neither is an abbreviation you have to decode later. `loc` means "find this label", `iloc` means "find this position".

```python
nums = pd.Series([5, 10, 15, 20, 25])
nums.loc[3]   # displayed as np.int64(20): the value with the NumPy type it is stored as. print() shows just 20.
```

```python
print(nums.loc[3])
```

```python
# the index can be replaced outright:
nums.index = ["a", "b", "c", "d", "e"]
nums
```

The label `3` no longer exists, so `nums.loc[3]` raises a `KeyError`, exactly as a dictionary does for a missing key. **The next cell is supposed to fail.** Read the last line of the traceback, then carry on: to reach the same value by label you now need `"d"`, and to reach it by position you use `.iloc`.

```python
# Asking for the old numeric label now raises a KeyError:
nums.loc[3]
```

```python
# By position instead: the fourth item, whatever its label happens to be.
nums.iloc[3]
```

So there are two ways to reach the same value: by **position** with `.iloc`, or by **label** with `.loc`. The two coincide only while the index is the default `0, 1, 2, ...`, which is why they looked interchangeable a moment ago.

---

### ✏️ Try it Out: Series Practice

Fill in the cell below, in order. Step 4 changes the index labels, so the two lookups have to happen before it.

1. Build a Series called `fruits` from the values `['apple', 'banana', 'cherry', 'date']` with the index `['a', 'b', 'c', 'd']`.
2. Use `.loc` to put the value at label `"c"` into `fruit_by_label`.
3. Use `.iloc` to put the **third** item into `fruit_by_position`.
4. Relabel the index to `['A1', 'B1', 'C1', 'D1']`, then display `fruits`.

```python
fruits = ...

fruit_by_label = ...
fruit_by_position = ...

# Relabel the index here, then display the Series.
fruits
```

```python
grader.check("ca05_series")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 3. DataFrame

A `DataFrame` is a two-dimensional table: a collection of `Series` that share one index. Each Series is a column, its name is the column name, and the shared row labels form the index.

### 3.1 Create a DataFrame from a `dict` object

```python
data = {'a': [1,2,3,4,500], 'b':[500, 700, 900, 1000, 500]} #this is a dictionary
df1 = pd.DataFrame(data)
df1
```

The keys became the column headings, and a default numerical index was generated.

---

### ✏️ Try it Out: build a DataFrame from a dictionary

Build a DataFrame called `cities_df` from a dictionary with two keys, `city` and `population`, each holding a list of three values. Use any three cities you like. Display it and check that the keys became the column headings.

```python
city_data = {
    ...
}

cities_df = ...
cities_df
```

```python
grader.check("ca05_dict_df")
```

### 3.2 Create a DataFrame by reading a CSV file

The function `.read_csv()` takes a file path as a string. A file path tells the computer where to find a file. The path can be **relative**, meaning it is read relative to the location of this notebook. If you move this notebook to a different folder, a relative path will no longer work.

If the data file sat in the same folder as this notebook, the path would be just `apartment_cost_list_20250217.csv`. Here it sits two folders down, so the relative path is `data/raw/apartment_cost_list_20250217.csv`.

```python
# https://data.cityofnewyork.us/Housing-Development/apartment-cost-list/wnmh-ssrd/about_data

df = pd.read_csv('data/raw/apartment_cost_list_20250217.csv')  # variable df now points to my dataframe
df
```

The difference from a `Series` is visible at once: several columns, each with a label. The `DataFrame` structure consists of:
- An index, with index labels (here the labels are just `0`, `1`, ..., `19595`)
- Columns, with column labels (here 'Job #', 'Borough', 'House #', 'Street Name', ....)
- And the data, which are the values in each row.

### 3.3 Create a DataFrame from an API response

The CSV above is a **snapshot**: someone exported the file on 2025-02-17 and it has not changed since. NYC Open Data also publishes that same dataset through an **API** (Application Programming Interface): a URL you can request records from directly, so you get the current version without downloading anything by hand.

This portal runs on **SODA**, the Socrata Open Data API, which is what most city open data portals use. The [Apartment Cost List](https://data.cityofnewyork.us/Housing-Development/apartment-cost-list/wnmh-ssrd/about_data) page has an "API" button that shows you the endpoint below.

#### Reading the request URL

```
https://data.cityofnewyork.us/resource/wnmh-ssrd.json
```

| Piece | Value | What it means |
|---|---|---|
| host | `data.cityofnewyork.us` | the portal serving the data |
| path | `/resource/` | SODA's endpoint for dataset records |
| dataset id | `wnmh-ssrd` | Socrata's "four-by-four" ID, four characters, a hyphen, four more. It is the last part of the dataset's own web address |
| format | `.json` | the format you want back (`.csv` and `.geojson` also work) |

#### The `params` argument

```python
requests.get(url, params={'$limit': 1_000})
```

`requests` turns the `params` dictionary into a **query string** and appends it to the URL. Print `response.url` after the request and you will see what was actually sent:

```
https://data.cityofnewyork.us/resource/wnmh-ssrd.json?%24limit=1000
```

`%24` is the escaped form of `$`, which the server reads back as `$`. You could type that URL out yourself, but passing a dictionary means you never have to work out those escapes for spaces, quotes or `&` characters.

`$limit` is one of SODA's query parameters. The `$` is part of the parameter's name, not Python syntax. These are the ones worth knowing:

| Parameter | What it does |
|---|---|
| `$limit` | how many records to return (1000 by default; some endpoints cap it at 50000) |
| `$offset` | how many records to skip, so you can page through a large dataset |
| `$select` | which columns to return |
| `$where` | filter rows, for example `"borough='BRONX'"` |
| `$order` | sort rows, for example `"job__ DESC"` |

#### The response object

`requests.get()` hands back a `Response` object, not the data itself. Three pieces of it matter here:

| Expression | What it gives you |
|---|---|
| `response.status_code` | the HTTP code: `200` means success, `404` not found, `429` too many requests |
| `response.ok` | `True` when the status code is below `400` |
| `response.json()` | parses the JSON body into Python objects |

The `if response.ok:` guard is why we read the body only when the server said the request worked. A failed request still returns JSON, but it is an error message rather than records, so without the guard you would quietly build a DataFrame out of the error and only notice much later.

```python
import requests

response = requests.get(
    'https://data.cityofnewyork.us/resource/wnmh-ssrd.json',
    params={'$limit': 1_000}  # This asks the API to return up to 1,000 records.
)

if response.ok:
    payload = response.json()
else:
    print(f'Request was not successful and returned code: {response.status_code}.')
    payload = None
```

#### What `.json()` gave you

For this endpoint, `response.json()` returns a **list of dictionaries**: one dictionary per record, with the field names as keys.

```python
print(type(payload))
print(len(payload), 'records')

payload[0]  # the first record, as a dictionary
```

`pd.DataFrame()` accepts that list directly: each dictionary becomes a row and each key becomes a column. If a record is missing a key, pandas fills that cell with `NaN`.

```python
# Create the DataFrame with the resulting payload:

apartments_df = pd.DataFrame(payload)
apartments_df.head(3)
```

The API and the CSV carry the same data under different labels. The API uses lowercase, underscored field names (`job__`, `street_name`), while the CSV keeps the portal's display names (`Job #`, `Street Name`). The API also drops fields a record has no value for, which is why `apartments_df` shows `NaN` in columns such as `curb_cut`. Watch for both whenever you move a notebook from a downloaded file to a live request.

### ✏️ Your Turn: read the request back

No check on this one: it depends on a live request, and the last step is a reading exercise rather than a value.

1. Send the request again with `$limit` set to `10`, then print `len(payload)` to confirm the API respected it.
2. Print `payload[0]` and pick out three field names.
3. Compare `apartments_df.columns` with `df.columns` and name the CSV column each of your three API fields corresponds to.

```python
# Your code here
```

<details>
<summary>Check your approach</summary>

```python
response = requests.get(
    'https://data.cityofnewyork.us/resource/wnmh-ssrd.json',
    params={'$limit': 10}
)
payload = response.json()

print(len(payload), 'records')
print(payload[0])

print(list(apartments_df.columns))
print(list(df.columns))
```

`job__` is `Job #`, `street_name` is `Street Name`, `bin__` is `Bin #`. The API lowercases the
name and replaces every space and punctuation mark with an underscore, so `Job #` becomes `job__`
and `Community - Board` becomes `community___board`.

</details>

### 3.4 Inspecting the `DataFrame`

`.shape` is an attribute rather than a method, so it takes no parentheses. It returns `(rows, columns)`:

```python
df.shape
```

`.head()` shows the first five rows.

```python
df.head() # first 5 rows by default
```

`.sample()` returns a random sample of rows; the argument is how many.

```python
df.sample(3)
```

`.describe()` gives summary statistics, but only for the numeric columns, so 7 of the 18 appear. Read it with care: `Job #`, `Block`, `Lot` and `Bin #` are identifiers, and the mean of an identifier is a number with no meaning.

```python
df.describe()
```

`.dtypes` lists the type pandas chose for each column. A column is numeric only if every value in it parsed as a number: `Initial Cost` came in as `object`, the label this version of pandas gives text, because of its `$` signs, and any arithmetic on it will have to wait until the text is converted. On pandas 3 the same column is labelled `str`; the meaning is identical.

```python
df.dtypes
```

`.info()` is the fuller summary: each column's type and how many non-null values it holds, then a count of columns by type and the memory used. `Curb Cut` holds 16 values in 19,596 rows, which `.dtypes` alone would never have told you.

```python
df.info()
```

#### index and columns

`.index` returns the row labels.

```python
df.index
## A RangeIndex, the default: pandas stores only the start, stop and step,
## so this is 19596 labels, 0 through 19595, without listing them.
```

`.columns` returns the column names, in order:

```python
df.columns
```

---

### ✏️ Try it Out: inspect the DataFrame

Answer three questions about `df` using the methods above rather than counting by hand. Read the numbers off the output of the cells you just ran, or compute them.

1. `n_rows`: how many rows does `df` have?
2. `n_numeric`: how many columns did pandas read as numbers rather than text? The last line of
   `.info()` states it; to compute it, `df.select_dtypes('number')` keeps only the numeric columns.
3. `first_three`: the first three column names, as a list.

```python
n_rows = ...
n_numeric = ...
first_three = ...

print('rows:', n_rows)
print('numeric columns:', n_numeric)
print('first three columns:', first_three)
```

```python
grader.check("ca05_inspect")
```

### 3.5 Selecting rows and columns: `.loc` and `.iloc`

Section 2.2 used `.loc` and `.iloc` on a Series, which has one axis. A DataFrame has two, so both
take two arguments: **rows first, then columns**.

| | rows | columns |
|---|---|---|
| `.loc[rows, cols]` | index **labels** | column **names** |
| `.iloc[rows, cols]` | row **positions** | column **positions** |

`df` still carries its default `RangeIndex`, so its row labels are `0, 1, 2, ...` and look exactly
like positions. They are not the same thing. The advanced part below shows where they come apart.

```python
# One whole row, by its index label. A single row comes back as a Series.
df.loc[0]
```

```python
# One value: the row labelled 0, the column named 'Borough'
df.loc[0, 'Borough']
```

```python
# A block: rows labelled 0 through 4, two named columns
df.loc[0:4, ['Borough', 'Street Name']]
```

Five rows, not four. **`.loc` slices include the last label.** This is the one place pandas parts
company with every other slice in Python, and it is worth remembering: `df.loc[0:4]` is five rows,
`my_list[0:4]` is four items.

```python
# The same block by position. Borough is column 1, Street Name is column 3.
df.iloc[0:5, [1, 3]]
```

`.iloc` slices behave the way Python slices always do: `0:5` stops **before** position 5, so this
is the same five rows. To ask for every row of one column, put a bare `:` in the row slot.

```python
print(df['Borough'].head(3).tolist())          # the shorthand you will use most
print(df.loc[:, 'Borough'].head(3).tolist())   # the same column, spelled out
```

#### Advanced 1: choosing rows with a condition

Instead of labels, `.loc` accepts a column of `True` and `False`, one per row, and keeps the `True`
ones. You build that column by comparing a column to a value: the comparison is vectorized, exactly
like the arithmetic in section 4 below.

```python
is_manhattan = df['Borough'] == 'MANHATTAN'   # one True or False per row
print(is_manhattan.head(3).tolist())
print('rows that are True:', is_manhattan.sum())

manhattan = df.loc[is_manhattan, ['Borough', 'Street Name', 'Initial Cost']]
manhattan.head(3)
```

Look at the index on the left: it starts at **3395**, not at 0. Filtering keeps each row's original
label, so labels and positions have now come apart. This is where mixing the two silently gives you
the wrong row.

```python
print('first row by position:', manhattan.iloc[0]['Street Name'])
print('row with label 3395: ', manhattan.loc[3395, 'Street Name'])

# There is no row labelled 0 left in manhattan, so manhattan.loc[0] would raise a KeyError.
```

#### Advanced 2: writing values, and the trap sitting next to it

To change a value, do the lookup and the assignment in **one** `.loc`. Reaching for the column
first and the row second is called **chained indexing**: the first bracket hands you a new object,
and the second one writes into that, which may not be the DataFrame you started with.

```python
print('pandas', pd.__version__)

sample = df.head(5).copy()   # .copy() so we never write into df itself
sample.loc[0, 'Borough'] = 'Bronx'          # one operation: the form to use
print('after .loc :', sample.loc[0, 'Borough'])

sample2 = df.head(5).copy()
sample2['Borough'][0] = 'Bronx'             # chained: two operations
print('after chain:', sample2.loc[0, 'Borough'])
```

The first form is the one to use: a single `.loc`, one operation, and the value changes.

The second is **chained indexing**, and what it does depends on which pandas you are running. On pandas 2.3, the version on the hub, it appears to work: the value changes to `Bronx`, and a `FutureWarning` announces that this will stop working in pandas 3. On pandas 3 it does stop working: the same line prints a `ChainedAssignmentError` warning and the value stays `BRONX`, because the edit went into a temporary copy that was then discarded. Code that leans on chained indexing is therefore correct on one version and silently wrong on the next.

That is the whole reason for the rule: **when you assign, put the row and the column inside a single `.loc[...]`.**

---

### ✏️ Try it Out: pull pieces out of `df`

1. `row_9_borough`: the borough recorded in the row **labelled** 9.
2. `first_five_streets`: the first five street names, as a list. (`list(...)` turns a Series into one.)
3. `queens`: every row whose `Borough` is `'QUEENS'`, keeping all the columns.

```python
row_9_borough = ...
first_five_streets = ...
queens = ...

print(row_9_borough)
print(first_five_streets)
print(queens.shape)
```

```python
grader.check("ca05_df_index")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 4. Vectorized Operations

A **vectorized** operation applies to an entire column at once instead of to one value at a time. Most operations on a `Series` or `DataFrame` work this way: the loop still happens, but inside compiled NumPy code rather than in your Python, which is why the same task runs many times faster on a large table and takes one line rather than four.

#### Example

You have a list of prices for grocery items, and you want to apply a **10% discount** to all of them. If you only know `for` loops, you might do something like this:

```python

prices = [10, 20, 30, 40]
discounted = []

for price in prices:
    discounted.append(price * 0.9)

discounted  # [9.0, 18.0, 27.0, 36.0]

```

With **pandas vectorization**, you can do this in **one line**:


```python
pd_prices = pd.Series([10, 20, 30, 40])
pd_discounted = pd_prices * 0.9

pd_discounted
```

The same numbers, with no loop written. The arithmetic operators all behave this way:

```python
sample = pd.Series([100, 200, 300, 7000000, 459023])
dash_ = '-'*20

print(sample + 5)   # Add 5 to each element
print(dash_)
print(sample * 2)   # Multiply each element by 2
print(dash_)
print(sample / 4)   # Divide each element by 4
print(dash_)
print(sample ** 2)  # Square each element
```

### Accessors: `.str`, `.dt`, `.cat`

Some methods only make sense for one kind of data. Lowercasing is meaningful for text but not for a date; asking for the year is meaningful for a date but not for text. Rather than crowd every method onto every Series, pandas puts them behind an **accessor**: a named doorway you open before calling the method.

```python
series.accessor.method()
```

| Accessor | Opens on a column of | Example |
|---|---|---|
| `.str` | text | `df['Street Name'].str.title()` |
| `.dt` | dates and times | `permits['issued'].dt.year` |
| `.cat` | categories | `df['Borough'].astype('category').cat.categories` |

Two things worth remembering:

- The accessor is not optional. `series.lower()` raises an `AttributeError`; `series.str.lower()` works.
- An accessor only exists when the column has the matching data type. `.dt` on a date column that pandas read as text will fail until you convert it with `pd.to_datetime()`, which is why `.dtypes` is worth checking first.

The rest of this section uses `.str`. You will meet `.dt` when we start working with dates.

### The `.str` accessor: vectorized methods mirroring Python string methods

> (Example from Python Data Science Handbook by Jake VanderPlas)

Nearly all of Python's built-in string methods have a vectorized twin behind `.str`, and they apply to the whole column at once. The full list is long; these are the ones this course reaches for:

| `.str` method | what you use it for |
|---|---|
| `len()` | length of each string (spot blanks or odd entries) |
| `lower()` / `upper()` | normalize case before comparing or grouping |
| `strip()` | trim leading/trailing spaces from messy text |
| `capitalize()` | capitalize the first letter |
| `startswith()` / `endswith()` | test a prefix or suffix (e.g. a code or label) |
| `split()` | break a string into parts (e.g. "City, State") |
| `zfill()` | pad numbers with leading zeros (zip / FIPS codes) |
| `isnumeric()` | check whether the text is all digits (validation) |

Their return values differ: `lower()` returns a Series of strings, `len()` a Series of integers, `startswith()` a Series of `True` and `False`, and `split()` a Series of lists. Start with the simplest:

```python

monte = pd.Series(['Graham Chapman', 'John Cleese', 'Terry Gilliam',
                   'Eric Idle', 'Terry Jones', 'Michael Palin'])
```

```python
monte.str.lower()
```

#### The same answer, written both ways

`.str.len()` gives the number of characters in each name. Here it is next to the `for` loop that
does the same job, so you can see what the accessor is doing for you.

```python
# The for loop version: somewhere to collect, then one pass over the names
lengths_loop = []
for name in monte:
    lengths_loop.append(len(name))

# The .str version: the same answer, in one line
lengths_str = monte.str.len()

print('loop:', lengths_loop)
print('.str:', lengths_str.tolist())
```

Same numbers. The difference is that `lengths_str` is still a **Series**: it kept `monte`'s index,
so each length stays attached to the name it came from, and you can feed it straight back into
pandas. `lengths_loop` is a plain list that has forgotten where the values came from.

---

### ✏️ Your Turn: string methods without a loop

Using the `monte` Series above, and no `for` loop. Each answer is one line.

1. `starts_with_t`: `True` or `False` for each name, depending on whether it starts with `'T'`.
2. `first_names`: the first name of each person. `.str.split()` splits each name on the space, and `.str.get(0)` takes the first piece of each split.

```python
starts_with_t = ...
first_names = ...

print(starts_with_t.tolist())
print(first_names.tolist())
```

```python
grader.check("ca05_str")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Summary

| Concept | Description |
|---------|-------------|
| **Series** | 1D labeled array, like a column |
| **DataFrame** | 2D labeled table, collection of Series |
| **`.loc[]`** | Location by label: `df.loc[row_label, column_name]` |
| **`.iloc[]`** | Integer location: `df.iloc[row_position, column_position]` |
| **`pd.read_csv()`** | Build a DataFrame from a CSV file path |
| **`pd.DataFrame(payload)`** | Build a DataFrame from a list of dictionaries, such as an API response |
| **`.head()`, `.sample()`** | View the first rows, or a random selection |
| **`.shape`** | (rows, columns) |
| **`.dtypes`** | Data type of each column |
| **`.info()`** | Column types, non-null counts, and the count of columns per type |
| **`.select_dtypes('number')`** | Only the numeric columns |
| **`.describe()`** | Summary statistics |
| **`.index`, `.columns`** | Row labels and column names |
| **Accessors (`.str`, `.dt`, `.cat`)** | Method namespaces for one data type, such as `.str.lower()` |

Tests for ca05_df_index (points: 0)
```
>>> assert row_9_borough == df.loc[9, 'Borough'], f"The row labelled 9 is in {df.loc[9, 'Borough']}. df.loc[9, 'Borough'] asks for one row and one column."
```
```
>>> assert first_five_streets == list(df.loc[0:4, 'Street Name']), f"Expected {list(df.loc[0:4, 'Street Name'])}. Remember .loc[0:4] is five rows because it includes the last label, while .iloc[0:5] is the same five by position."
```
```
>>> assert set(queens['Borough'].unique()) == {'QUEENS'}, "queens should hold only Queens rows. Build the condition with df['Borough'] == 'QUEENS' and pass it to .loc."
```
```
>>> assert queens.shape == df.loc[df['Borough'] == 'QUEENS'].shape, f"Expected {df.loc[df['Borough'] == 'QUEENS'].shape} (rows, columns), got {queens.shape}. Keep every column: pass only the row condition to .loc."
```


Tests for ca05_dict_df (points: 0)
```
>>> assert isinstance(cities_df, pd.DataFrame), 'cities_df should be a DataFrame. Pass your dictionary to pd.DataFrame(...).'
```
```
>>> assert list(cities_df.columns) == ['city', 'population'], f"The dictionary keys become the column names, so expect ['city', 'population'], not {list(cities_df.columns)}."
```
```
>>> assert cities_df.shape[0] == 3, f'Three cities means three rows. Yours has {cities_df.shape[0]}.'
```


Tests for ca05_inspect (points: 0)
```
>>> assert n_rows == df.shape[0], f'df has {df.shape[0]} rows. .shape returns (rows, columns), so the row count is the first item.'
```
```
>>> assert n_numeric == df.select_dtypes('number').shape[1], f"{df.select_dtypes('number').shape[1]} columns came in as numbers. .dtypes reports them as int64 and the text columns as object; df.select_dtypes('number').shape[1] counts the numeric ones."
```
```
>>> assert first_three == list(df.columns[:3]), f'Expected {list(df.columns[:3])}. df.columns holds the names in order, and it can be sliced.'
```


Tests for ca05_series (points: 0)
```
>>> assert isinstance(fruits, pd.Series), 'fruits should be a Series. Build it with pd.Series(...).'
```
```
>>> assert list(fruits) == ['apple', 'banana', 'cherry', 'date'], 'Keep the four values, in the order they are given.'
```
```
>>> assert list(fruits.index) == ['A1', 'B1', 'C1', 'D1'], 'Step 4 asks for the index labels A1, B1, C1, D1. Assign a new list to fruits.index.'
```
```
>>> assert fruit_by_label == 'cherry', "fruits.loc['c'] reads the value at the label 'c', which is 'cherry'. If you relabelled the index first, .loc['c'] no longer exists."
```
```
>>> assert fruit_by_position == 'cherry', 'The third item sits at position 2, because .iloc counts from 0.'
```


Tests for ca05_str (points: 0)
```
>>> assert list(starts_with_t) == [False, False, True, False, True, False], "monte.str.startswith('T') returns one True or False per name. Only the two Terrys start with a capital T."
```
```
>>> assert list(first_names) == ['Graham', 'John', 'Terry', 'Eric', 'Terry', 'Michael'], 'monte.str.split() gives a list per name; .str.get(0) takes the first item of each list.'
```
