<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# CP101: Lab 2 - Computational Thinking for Urban Planners

---

## What is Computational Thinking?

Before we write a single line of code, we need to learn how to **think computationally**. This doesn't mean becoming a robot—it means developing a systematic approach to problem-solving that computers can help us execute.

Computational thinking has four key components:

| Component | What it means | Urban Planning Example |
|-----------|--------------|------------------------|
| **Decomposition** | Breaking a complex problem into smaller, manageable parts | Breaking "improve transit access" into: identify underserved areas, map existing routes, analyze demographics, propose new stops |
| **Pattern Recognition** | Finding similarities or trends in data | Noticing that traffic accidents cluster at intersections without protected left turns |
| **Abstraction** | Focusing on important information while ignoring irrelevant details | Representing a city as a network of nodes (intersections) and edges (streets) rather than every building |
| **Algorithm Design** | Creating step-by-step instructions to solve a problem | A process for determining which potholes to fix first based on severity, traffic volume, and budget |

---

```python
# Initialize Otter
import otter
grader = otter.Notebook("CP101_LAB02.ipynb")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Activity 1: Decomposition - The BART Accessibility Problem

### The Scenario

The Metropolitan Transportation Commission (MTC) has asked you: **"Which Bay Area neighborhoods have the worst access to BART stations?"**

This is a big, vague question. A computer can't answer it directly. Neither can we, until we break it down.

### **Your Task: Decompose This Problem**

In the cell below, break down the main question into smaller sub-questions. I've started with a few examples.

**Main Question:** Which Bay Area neighborhoods have the worst access to BART stations?

**Sub-questions:**
1. What do we mean by "neighborhood"? (Census tracts? Zip codes? Named neighborhoods?)
2. How do we define "access"? (Distance? Travel time? Number of stations within X miles?)
3. Where are all the BART stations located?
4. _[Add your own sub-question here]_
5. _[Add your own sub-question here]_
6. _[Add your own sub-question here]_
7. _[Add your own sub-question here]_

---

### Your Decomposition (Double-click to edit)

4. -
5. -
6. -
7. -

<!-- END QUESTION -->

### Reflection

Notice how the original question was impossible to answer directly, but each sub-question is much more concrete. This is the power of decomposition!

---

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Activity 2: Pattern Recognition - SF 311 Service Requests

San Francisco's 311 system receives thousands of service requests every day. Below is a sample of real data from the SF Open Data portal.

### The Data

| Request ID | Category | Neighborhood | Day of Week | Response Time (hours) |
|------------|----------|--------------|-------------|----------------------|
| 001 | Street Cleaning | Mission | Monday | 24 |
| 002 | Street Cleaning | Mission | Monday | 26 |
| 003 | Graffiti | Castro | Tuesday | 48 |
| 004 | Pothole | Sunset | Monday | 120 |
| 005 | Street Cleaning | Mission | Wednesday | 22 |
| 006 | Graffiti | Tenderloin | Monday | 72 |
| 007 | Pothole | Richmond | Tuesday | 96 |
| 008 | Street Cleaning | Tenderloin | Monday | 36 |
| 009 | Graffiti | Mission | Thursday | 52 |
| 010 | Pothole | Sunset | Friday | 144 |

### Your Task: Identify Patterns

Look at this data and identify at least **three patterns**. A pattern could be about:
- Which categories are most common
- Which neighborhoods appear frequently
- How response times vary by category
- Any other trends you notice

First, record what the data says. Assign each variable below to a **string** (copy the spelling from the table):

- `most_common_category`: the category with the most requests
- `most_frequent_neighborhood`: the neighborhood that appears most often
- `slowest_category`: the category with the longest response times
- `busiest_day`: the day of the week with the most requests

```python
most_common_category = ...
most_frequent_neighborhood = ...
slowest_category = ...
busiest_day = ...
```

```python
grader.check("q2")
```

<!-- BEGIN QUESTION -->

---

### Your Patterns (Double-click to edit)

**Pattern 1:**

**Pattern 2:**

**Pattern 3:**

<!-- END QUESTION -->

### Why This Matters

When we learn Python, we'll be able to analyze thousands of 311 requests instead of just 10. But the skill of recognizing patterns? That's something YOU bring to the table. The computer just helps you see patterns at scale.

---

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Activity 3: Abstraction - Mapping Oakland's Street Network

### The Problem

Imagine you're analyzing Oakland's street network to find the best route for a new bus line. Oakland has:
- ~2,000 miles of streets
- ~15,000 intersections
- Varying speed limits, one-way streets, turn restrictions
- Thousands of buildings, trees, fire hydrants, etc.

Should we include ALL of this information in our analysis?

### The Power of Abstraction

**Abstraction** means keeping only what matters for our specific question and ignoring the rest.

For bus route planning, we might abstract Oakland's streets into:
- **Nodes**: Intersections and bus stops
- **Edges**: Street segments connecting nodes
- **Attributes**: Travel time, population density along segment

We deliberately **ignore**: building colors, tree species, fire hydrant locations—these don't affect bus routing.

### Your Task

For each scenario below, list what information you would **KEEP** and what you would **IGNORE**:

**Scenario A: Identifying neighborhoods vulnerable to sea level rise**
- Keep:
- Ignore:

**Scenario B: Finding the best location for a new library**
- Keep:
- Ignore:

**Scenario C: Analyzing pedestrian safety at intersections**
- Keep:
- Ignore:

---

### Your Abstractions (Double-click to edit)

**Scenario A: Sea Level Rise**
- Keep:
- Ignore:

**Scenario B: Library Location**
- Keep:
- Ignore:

**Scenario C: Pedestrian Safety**
- Keep:
- Ignore:

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Activity 4: Algorithm Design - The Pothole Prioritization Problem

### The Scenario

You work for the Oakland Department of Transportation. You have:
- 500 reported potholes
- Budget to fix 100 potholes this month
- Data on each pothole: size, location, traffic volume, time since reported

### What is an Algorithm?

An **algorithm** is a step-by-step procedure for solving a problem. It's like a recipe: if you follow the steps exactly, you'll get the same result every time.

Algorithms must be:
- **Unambiguous**: Each step is clear
- **Finite**: It eventually ends
- **Effective**: Each step can actually be performed

### Example Algorithm: Making a PB&J Sandwich

```
1. Get two slices of bread
2. Get peanut butter jar and jelly jar
3. Get a knife
4. Open peanut butter jar
5. Use knife to scoop peanut butter
6. Spread peanut butter on one slice of bread
7. Open jelly jar
8. Use knife to scoop jelly
9. Spread jelly on other slice of bread
10. Press the two slices together, spreads facing inward
11. Done!
```

### Your Task: Design a Pothole Prioritization Algorithm

Write an algorithm (in plain English, like the PB&J example) that determines which 100 potholes to fix. Your algorithm should be specific enough that anyone following it would make the same decisions.

Think about:
- What criteria matter most?
- How do you handle ties?
- What order do you consider factors?

---

### Your Algorithm (Double-click to edit)

**Pothole Prioritization Algorithm**

```
1. [Your first step]
2. [Your second step]
3. ...
```

<!-- END QUESTION -->

### Discussion Questions

1. Did you and your neighbor come up with the same algorithm? Why or why not?

2. What values or priorities are embedded in your algorithm? (e.g., Does it favor high-traffic areas? Recently reported issues? Equity considerations?)

3. Could a computer execute your algorithm? What additional information would it need?

---

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Activity 5: Putting It All Together - The Complete Workflow

### The Big Picture

In this course, you'll learn to solve urban planning problems using this workflow:

```
Real-World Problem
       ↓
   [DECOMPOSITION]
       ↓
Smaller Sub-Problems
       ↓
   [PATTERN RECOGNITION]
       ↓
Similarities / Reusable Structures
       ↓
   [ABSTRACTION]
       ↓
Essential Details + Model (ignore irrelevant detail)
       ↓
   [ALGORITHM DESIGN]
       ↓
Step-by-Step Solution
       ↓
   [EVALUATION / ITERATION]
       ↓
Test, Debug, Refine (repeat as needed)

```

### Final Exercise

Choose ONE of these urban planning questions and apply all four computational thinking skills:

**Option A:** "How can we reduce traffic congestion on the Bay Bridge?"

**Option B:** "Where should Oakland build more affordable housing?"

**Option C:** "How can San Francisco improve response times for 311 requests?"

For your chosen question, write:
1. **Decomposition:** Break it into 4-5 sub-questions
2. **Abstraction:** What data would you need? What would you ignore?
3. **Algorithm:** Sketch a high-level algorithm (5-10 steps)
4. **Pattern Recognition:** What patterns might you look for in the results?

---

### Your Complete Analysis (Double-click to edit)

**Question I'm analyzing:** [Write which option you chose]


#### 1. Decomposition (Sub-questions):

-
-
-
-

#### 2. Abstraction (Data needed vs. ignored):

- Data to collect:
- Data to ignore:


#### 3. Pattern Recognition (What patterns to look for):
-
-
-

#### 4. Algorithm (High-level steps):

```
1.
2.
3.
4.
5.
```


<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Reflection: Why Computational Thinking Before Coding?

You might be eager to start writing Python code. So why did we spend this whole lab on computational thinking?

Here's the secret: **The hard part of programming isn't the code Syntax, it's the thinking.**

Professional programmers spend most of their time:
- Understanding the problem
- Breaking it into pieces
- Designing solutions
- Debugging their thinking

The actual typing? That's maybe 20% of the job.

When we start Python next week, you'll see how these computational thinking skills translate directly into code. Loops implement algorithms. Functions encapsulate sub-problems. Data structures are abstractions. Pattern recognition becomes data analysis.

**The better you are at computational thinking, the better programmer you'll become.**

**Next week:** We dive into Python!

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload your submission zip to **LAB02 - Computational Thinking** on [Gradescope](https://www.gradescope.com/).

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q2 (points: 1)
```
>>> assert most_common_category.strip().title() == 'Street Cleaning'
```
```
>>> assert most_frequent_neighborhood.strip().title() == 'Mission'
```
```
>>> assert slowest_category.strip().title() == 'Pothole'
```
```
>>> assert busiest_day.strip().title() == 'Monday'
```
