<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# CP101: Lab 1 - JupyterHub, Jupyter Notebooks, and Submitting Your Work

### CP101: Introduction to Urban Data Analytics

**Lab 1 is two notebooks, and we do both today.**

**This one** gets you set up: DataHub, how notebooks work, and one full trip through submitting to Gradescope.

**My Places** (`CP101_MY_PLACES.ipynb`, in this same folder) starts the dataset you will build on for the rest of the semester.

This is the lab where we set up the machinery you will use every week from here on. There is no urban data analysis in it. That starts next week.

**We do Part 3 together, in the room.** It is the one part where things go wrong for someone every semester, and the whole point of doing it now, on a lab that carries no stakes, is that your GSI is standing right there when it does.

```python
# Initialize Otter
import otter
grader = otter.Notebook("CP101_LAB01.ipynb")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Learning Objectives

By the end of this lab you will be able to:

1. Log in to DataHub and describe what it is and where your files live.
2. Find your way around the JupyterLab interface: the file browser, the launcher, and the menu bar.
3. Tell a code cell from a markdown cell, and create, run, and delete both.
4. Explain what the kernel is and why the order you run cells in matters.
5. Generate a submission zip and upload it to Gradescope.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## How this lab works

Read the sections in order and do each ✏️ **Try it Out** as you reach it. The activities build on each other, and **the cells you add are your submission**, so leave them in the notebook when you are done.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 1: JupyterHub and DataHub

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 1.1 What DataHub is, and why we use it

**DataHub** ([datahub.berkeley.edu](https://datahub.berkeley.edu)) is UC Berkeley's hosted **JupyterHub**: a shared computing environment that runs in the cloud and that you reach through a web browser.

You log in with your CalNet ID and you get a Linux machine with Python, JupyterLab, and every package this course needs already installed. There is nothing to download and nothing to configure.

That last point matters more than it sounds. In a course where everyone installs Python themselves, a meaningful share of the first month goes to fixing installations, and the students it hits hardest are the ones with the oldest laptops. On DataHub everyone gets the same environment, so when your code behaves differently from your neighbour's it is because of your code, not your machine.

**Anyone with a Berkeley email has access.** There is no separate signup.

> **A word on the word.** *JupyterHub* is the software that hands each user their own server. *DataHub* is Berkeley's installation of it. *JupyterLab* is the interface you see once you are in. People use the three names loosely; now you know which is which.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 1.2 What you see when you log in

After CalNet authentication, JupyterLab opens. Four regions, and you will use all of them today.

| Region | Where | What it is for |
|---|---|---|
| **Menu bar** | across the top | File, Edit, View, Run, Kernel, Settings |
| **Left sidebar** | far left icons | File browser, running kernels, table of contents, extensions |
| **Main work area** | the middle | Your open notebooks, in tabs |
| **Launcher** | a tab in the work area | Buttons to start a new notebook, terminal, or text file |

If the Launcher is not showing, click the blue **+** at the top of the file browser.

### The menus you will actually use

| Menu | What you will need it for |
|---|---|
| **File** | Save, Download, and the PDF export |
| **Edit** | Undo, and moving cells around |
| **View** | Turning the left sidebar and table of contents on and off |
| **Run** | Run All Cells, and Run All Above |
| **Kernel** | **Restart Kernel**, the single most useful troubleshooting step there is |

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 1.3 Where your files live

The file browser in the left sidebar shows your **home directory** on DataHub. It is yours, it is private, and it persists between sessions: close the tab, come back tomorrow, and your files are where you left them.

Course materials are distributed into folders, one per lab or assignment. A lab folder looks like this:

```text
CP101_LAB01.ipynb        <- the notebook you are reading
CP101_MY_PLACES.ipynb    <- the My Places notebook, which you do after this one
CP101_LAB02.ipynb        <- next week's lab, already here
sample_py_file.py        <- a plain Python script, for comparison in Section 2.1
template_myplaces.csv    <- the file My Places has you fill in
```

Labs that read data carry one more thing, a `data/` folder with a `raw/` folder inside it.
This lab has no data of its own, but you will meet that layout from Lab 3 onward, and it is
worth knowing what it means now.

### Why `data/raw/`

Nearly every lab in this course reads external data from a `data/raw/` folder, and the convention is worth naming because it is a professional habit rather than a Jupyter rule.

**`raw` means untouched.** Whatever is in there is the data as it came from the Census Bureau, the city portal, or the survey. You read from it and you never write back to it. Your cleaning happens in memory, in the notebook, where anyone can see exactly what you did to the data.

The reason is reproducibility. If you clean a file in place, the cleaning becomes invisible and permanent: nobody else can rerun your work from the source, including you in three weeks. Keeping the raw file untouched means your notebook is a complete, checkable account of how you got from the source to your result.

### Paths are relative

When a lab reads a file, it does so like this:

```python
pd.read_csv('data/raw/some_file.csv')
```

That is a **relative path**: it means "starting from the folder this notebook is in". Written that way, the code runs on your DataHub account, on your GSI's, and on the autograder, because all three have the notebook and its `data/` folder sitting together. An **absolute path** like `/home/jovyan/CP101/data/raw/some_file.csv` works only on the one machine it was typed on, which is why you will lose points for them later in the course.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## ✏️ Try it Out 1: find your way around

Nothing to write for this one. Just do it, so you know where things are.

1. Open the **file browser** (the folder icon at the top of the left sidebar).
2. Find `CP101_LAB01.ipynb`, this notebook, and `sample_py_file.py` next to it. You should also see `CP101_MY_PLACES.ipynb`, the other half of this lab.
3. **Open `sample_py_file.py`.** It opens in a plain text editor rather than a notebook view. Look at it for a moment: it is nothing but code and `#` comments, top to bottom, with no output saved anywhere. Then close its tab. That contrast is the subject of Section 2.1.
5. Click the **Table of Contents** icon in the left sidebar and use it to jump to Part 3 of this notebook, then back up here.

> **If you get lost in the file browser**, the folder name breadcrumbs just above the file list will take you back up.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 1.4 Saving, and what can go wrong!

JupyterLab autosaves every couple of minutes, but get in the habit of `Ctrl+S` (`Cmd+S` on Mac). You will need it in Part 3, because **the submission export reads the file on disk, not what is on your screen.** Unsaved edits are not in your zip.

### Back your work up

DataHub home directories are **not long-term storage**, and Berkeley says so explicitly. Files left untouched for 90 days get archived off to cold storage, and the hub is not a backup service.

So: at the end of the semester, and honestly at the end of each project, download what you care about and put it somewhere of your own. Google Drive, Dropbox, a GitHub repo, your laptop. Right-click any file in the file browser and choose **Download**.

This is not a hypothetical. People lose final project work this way every year.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 1.5 When something goes wrong

Work through these in order. The first one fixes most things.

| Symptom | What to try |
|---|---|
| A cell runs forever, or the notebook stops responding | **Kernel → Interrupt Kernel**, then **Kernel → Restart Kernel** |
| Weird errors about variables that should exist | **Kernel → Restart Kernel and Run All Cells**, see Section 2.4 |
| `ModuleNotFoundError` for a package the lab uses | Check you are on the course kernel, then ask; do not start pip-installing |
| The whole server is slow or will not start | **File → Hub Control Panel → Stop My Server**, then start it again |
| You cannot log in at all | Check your CalNet ID is active, then ask your GSI |

**Restarting the kernel does not delete anything.** It throws away the variables Python is holding in memory and gives you a clean slate; your notebook, your code, and your files are untouched. When in doubt, restart.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 2: Jupyter Notebooks

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 2.1 `.py` files and `.ipynb` files

Python code lives in one of two formats, and you just looked at one of each.

- **`.py` files** are plain-text scripts: just code, run top to bottom as a single program each time, with no output stored in the file. `sample_py_file.py` is one.
- **`.ipynb` files** ("IPython Notebook") are structured documents made of separate **cells**. Cells can be run individually and in any order, they share one live Python session, and each code cell's output is saved back into the file underneath the code that produced it. This notebook is one.

This course works almost entirely in notebooks, because mixing explanation, code, and results in one document is exactly what data analysis needs. When you get to writing reusable functions and scripts, `.py` files come back.

| | `.py` | `.ipynb` |
|---|---|---|
| Structure | one continuous file | separate cells |
| Execution | top to bottom, every time | any cell, any order |
| Output | printed to a terminal, not stored | saved into the file |
| Narration | `#` comments only | full markdown text cells |
| Good for | reusable programs, scripts | analysis, teaching, showing your work |

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 2.2 The two kinds of cell

**Code cells** hold Python. Run one and it executes, and anything it prints or returns appears directly underneath it.

**Markdown cells** hold formatted text: headings, lists, links, tables, images. Run one and it renders. Everything you have been reading in this lab is a markdown cell.

There is a third type, **raw**, which is passed through untouched. You will rarely see it and you will not need to make one.

The cell type is shown in the dropdown in the toolbar, and you can change it whenever you like. A cell that is showing you a wall of unformatted text with `#` symbols in front of it is a markdown cell you have not run yet.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 2.3 Command mode and edit mode

This is the piece that trips everyone up, so it comes before the shortcuts rather than after.

JupyterLab has two modes:

- **Edit mode**: there is a cursor inside the cell and typing types. Get there by pressing `Enter` or clicking into the cell.
- **Command mode**: the cell is selected as a whole and single keystrokes are commands. Get there by pressing `Esc` or clicking the cell's left margin. The border is highlighted and there is no blinking cursor.

**Every keyboard shortcut below is a command mode shortcut.** If pressing `B` types the letter b into your cell, you are in edit mode. Press `Esc` first.

| Action | Key | Note |
|---|---|---|
| Insert cell **below** | `B` | |
| Insert cell **above** | `A` | |
| Change to **markdown** | `M` | |
| Change to **code** | `Y` | |
| **Delete** cell | `D`, `D` | press it twice |
| **Undo** cell operation | `Z` | brings a deleted cell back |
| Run and move on | `Shift+Enter` | |
| Run and stay put | `Ctrl+Enter` | |

> `Z` in command mode undoes **cell** operations. `Ctrl+Z` in edit mode undoes your **typing**. After a stray `D`,`D` the one you want is `Z`.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## ✏️ Try it Out 2: add a markdown cell

1. Click this cell, then press `Esc` to be sure you are in command mode.
2. Press `B` to insert a new cell below this one.
3. Press `M` to make it a markdown cell.
4. Write one or two sentences, in your own words, on how a `.py` file differs from an `.ipynb` file. Use at least one piece of markdown formatting from Section 2.6, such as **bold** or a bulleted list.
5. Press `Shift+Enter` to render it.

**Leave the cell in the notebook.** It is part of what you hand in.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 2.4 Running cells, and the kernel

Behind your notebook is a **kernel**: a live Python session that stays running between cells. This is the single most important thing to understand about notebooks.

Because the kernel persists:

- A variable defined in one cell is available in every other cell.
- Cells can be run in **any order**, and the kernel remembers whatever you last told it.
- The `[1]`, `[2]` numbers in the left margin are the **execution count**: the order cells were actually run in, which is not necessarily the order they appear on the page.
- A cell showing `[*]` is still running.

### Why this bites people

Suppose you define `x = 5`, run the cell, then edit it to `x = 500` but do not rerun it. The kernel still holds `5`. Your notebook now says one thing on screen and another in memory, and every result below it is computed from a value that is not written anywhere.

The habit that prevents this: before you submit anything, **Kernel → Restart Kernel and Run All Cells**. If it runs top to bottom without error, your notebook actually says what it appears to say. If it does not, you have found a real problem and you found it before your GSI did.

You will see this instruction in every lab from here on. This is why.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## ✏️ Try it Out 3: add and run a code cell

1. In command mode, press `B` to add a cell below this one. Leave it as a code cell.
2. Type `print('Hello World')` and run it with `Shift+Enter`.
3. Run it again with `Ctrl+Enter` and notice the difference: one moves the selection on, the other stays put.
4. Watch the number in brackets go up each time you run it.

Then, to see the kernel's memory for yourself:

5. Add another code cell. In it, put `my_number = 5` on one line and `print(my_number)` on the next. Run it.
6. Now change `5` to `500`, but **do not run the cell**. Instead add one more cell below and put just `print(my_number)` in it. Run only that new cell.

It prints `5`, not `500`, because the kernel still holds what you last *ran*, not what is on screen. Now go back, run the edited cell, then run the print cell again.

**Leave all of these cells in the notebook.**

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## ✏️ Try it Out 4: delete a cell, and bring it back

Deleting is the operation people are most nervous about, so practise it where nothing matters.

1. Add a cell below this one with `B` and type anything into it.
2. Press `Esc`, then `D` twice. The cell is gone.
3. Press `Z`. It comes back, contents and all.

That is the whole trick. You are very unlikely to destroy anything you cannot recover in the next keystroke.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 2.5 Markdown, briefly

Markdown is a lightweight way of formatting text. You need a handful of it for writing up your answers.

| What you type | What you get |
|---|---|
| `# Heading 1` to `###### Heading 6` | headings, largest to smallest |
| `**bold**` | **bold** |
| `*italic*` | *italic* |
| `- item` | a bulleted list |
| `1. item` | a numbered list |
| `` `code` `` | `inline code` |
| `[text](https://example.com)` | a [link](https://example.com) |

For a block of code rather than a phrase, fence it with three backticks and name the language:

````text
```python
print("Hello, World!")
```
````

And a table, which is the one piece of markdown that looks worse than it is:

```text
| Column 1 | Column 2 |
|----------|----------|
| Row 1    | Data     |
```

> **Tip:** to see how any formatting in this lab was written, double-click the cell to open it, then press `Shift+Enter` to put it back. This is the fastest way to learn markdown, and you cannot break anything.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 2.6 How your work gets checked

Most labs in this course, including this one, are graded by a computer. Here is how that works, so nothing about it is a surprise.

### The checking happens after you submit

When you upload your zip to Gradescope, an **autograder** opens your notebook, runs it, and tests specific answers. A few moments later Gradescope shows you the result: which checks passed, which did not, and a message for each one that failed.

So you do not have to guess whether you got it right. You submit, and you are told.

When a check fails you do not get a bare `AssertionError`. You get a message written to tell you what went wrong:

```text
AssertionError: twenty_twenty_six should work out to 2026. If this failed, check the
arithmetic in your expression.
```

Read those messages. They are the fastest help you will get in this course, and they are written by the same people who wrote the question.

**You can resubmit** as many times as you like before the deadline, so a failed check is information, not a verdict. Fix it, export again, upload again.

### Some labs also let you check as you go

In later labs you will meet cells like `grader.check("q1")`, which run a question's tests right there in the notebook so you can check yourself before submitting. This lab does not have them, on purpose: today the point is to see the submit-and-get-graded loop from end to end.

### Hints you can reveal

Some exercises carry a hint folded out of the way, so it is there when you are stuck and not in your face when you are not. Click the triangle:

<details>
<summary>Hint: click this triangle</summary>

This is a revealed hint. In a real lab it would point you at the method you need, or name the mistake people usually make, without writing the answer out for you.

</details>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## ✏️ Try it Out 5: your first graded question

Everything above was practice you checked yourself. This one is different: it is **graded by a computer**, and you will see the result after you submit.

Assign `twenty_twenty_six` to an arithmetic expression that equals `2026`.

Do not just write the number. The point is to write an **expression**, something like a multiplication or an addition, that works out to 2026. There are many right answers.

There is no button here to check yourself. That is deliberate. When you upload your zip in Part 3, Gradescope runs a check on this cell and shows you whether it passed, which is the whole loop you are practising today.

```python
twenty_twenty_six = ...

print(twenty_twenty_six)
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 3: Submitting to Gradescope

**We do this part together in class.** Work through it at the same pace as everyone else and put your hand up the moment something does not match what is described here. Getting stuck now, with help in the room, is the entire purpose of this lab.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 3.1 What you submit, and what the export cell does

For most labs you submit a **zip file** that you generate from inside the notebook.

The export cell bundles your notebook, and any files the lab needs, into a single timestamped `.zip` and gives you a download link. That zip is what Gradescope reads.

Three things worth knowing before you run it:

1. **It reads the saved file.** Save first, with `Ctrl+S`. Unsaved edits are not in the zip.
2. **The filename has a timestamp in it**, something like `CP101_LAB01_2026_08_18T12_55_09.zip`. That is normal. Do not rename it.
3. **You can run it as many times as you like.** Each run makes a fresh zip. If you fix something after exporting, just export again and upload the new one.

Some assignments also ask for a **PDF**, for written answers a human reads rather than a test. Those labs will say so and will include a second export cell. This one does not.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 3.2 Do this now, with the room

**Step 1. Check your notebook runs.** Go to **Kernel → Restart Kernel and Run All Cells** and let it finish. It should run top to bottom without stopping. Your Try it Out cells should all still be in the notebook, and Try it Out 5 should print your answer.

**Step 2. Save.** `Ctrl+S`, or `Cmd+S` on a Mac. The export reads the saved file, so anything you have not saved will not be in your zip.

**Step 3. Run the export cell** at the very bottom of this notebook.

**Step 4. Download the zip.** Click the download link the cell prints. If no link appears, find the `.zip` in the file browser on the left, right-click it, and choose **Download**. It goes to your computer's Downloads folder.

**Step 5. Upload it.** Go to [Gradescope](https://www.gradescope.com/), open **LAB01 - JupyterHub, Jupyter Notebooks, and Submitting Your Work**, and upload the zip.

**Step 6. Watch it get graded.** This is the part worth waiting for. Gradescope runs the autograder on your notebook and, after a few moments, shows you the checks and whether they passed. That result is coming from the `twenty_twenty_six` cell you wrote in Try it Out 5.

If a check failed, read the message, fix the cell, save, export again, and upload again.

> **You can resubmit.** Gradescope keeps the most recent submission before the deadline, so a mistake now costs you nothing. If anything does not look like this description, ask.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 3.3 The export cell

The export cell is the **last cell in this notebook**, below the recap. Scroll to the bottom, run it, and it builds your submission zip.

Save first.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 3.4 If something goes wrong

| What you see | What to do |
|---|---|
| The cell errors with a name or import problem | **Kernel → Restart Kernel and Run All Cells**, then run the export cell again |
| No download link appears | Find the `.zip` in the file browser, right-click, **Download** |
| The zip is not in your Downloads folder | Check your browser's downloads list; some browsers ask where to put it |
| Gradescope rejects the file | Make sure you are uploading the `.zip`, not the `.ipynb` and not a folder |
| You uploaded, then changed something | Export again and upload again; the newest submission is the one that counts |
| Your work is missing from the zip | You did not save before exporting. Save, export again, upload again |

**If you are stuck for more than a minute, put your hand up.** That is what today is for.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Recap

| Thing | What to remember |
|---|---|
| DataHub | Berkeley's hosted JupyterHub; log in with CalNet, same environment for everyone |
| Your files | Persist between sessions, but are **not** long-term storage; back them up |
| `data/raw/` | The data as it arrived. Read from it, never write to it |
| Relative paths | `data/raw/file.csv`, not `/home/jovyan/...`, so your code runs anywhere |
| Command vs edit mode | `Esc` for command mode; every shortcut lives there |
| The kernel | One live Python session; it remembers what you **ran**, not what is on screen |
| Before submitting | **Kernel → Restart Kernel and Run All Cells**, then save |
| Submitting | Run the export cell, download the zip, upload to Gradescope, confirm |

### Next Steps

My Places has you build your own small dataset as a CSV file and practise submitting **that**, which is the other kind of upload you will need this semester.

Lab 2 leaves the machinery behind and starts on computational thinking. 

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload your submission zip to **LAB01 - JupyterHub, Jupyter Notebooks, and Submitting Your Work** on [Gradescope](https://www.gradescope.com/).

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1 (points: 1)
```
>>> assert twenty_twenty_six == 2026, 'twenty_twenty_six should work out to 2026. If this failed, check the arithmetic in your expression.'
```
