<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# CP101: Lab 4 - Python Basics, Part II

1.   Lists
2.   Conditionals
3.   Loops
4.   Boolean Expressions

```python
# Initialize Otter
import otter
grader = otter.Notebook("CP101_LAB04.ipynb")
```

Data Structures:

![data structures](images/data_structures.png)

Lets focus on Sequences for a moment:

Sequences are ordered collections that are iterable and indexable.

In python, lists, strings (as we saw last week), and tuples. Lists are mutable, strings and tuples are not.




<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 1. Lists

Lists are used to store multiple items in a single variable. They are one of the for main data structures that base data types(str, int, float and complex) in python can be stored in.

Lists are created by using square brackets.

```python
list1 = [1,2,3,4]
list2 = ["Dog", "Cat", "Mouse", "Froggy"]
list3 = [True, False, True, True]

print(list1)
print(list2)
print(list3)
```

Common list methods that will help us later on:

For adding elements:
*   **append():** Adds an element at the end of the list.
*   **extend():** Adds multiple elements to the end of the list.
*   **insert():** Adds an element at a specific position.
*   **clear():** removes all items.

For removing elements:
*   **pop():** removes the first instance of an element
*   **remove():** removes an element at a specific index OR the last element if no index argument is specified.

```python
# Let's first create an empty list
demo_list = []

# Let's add some stuff to the end of the list
demo_list.append("hello")
print(demo_list)

# Let's add multiple elements to the end
demo_list.extend([5, 12.13, "hello"])
print(demo_list)

# Let's insert something at a specific position
demo_list.insert(2, True)
print(demo_list)

# Oops I didn't mean to add a Boolean True, lets update our list with indexing
demo_list[2] = False
print(demo_list)

```

```python
# I found some stuff in my list that I dont want, let's remove them
demo_list.pop(1)
print(demo_list)

# I added "hello" one too many times, let's remove the first one
demo_list.remove("hello")
print(demo_list)
```

```python
# This list is full of nonsense, let's clear it
demo_list.clear()
print(demo_list)
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 2. Conditionals

Comparison Operators

| Operator | Meaning | Example |
|----------|---------|--------|
| `==` | Equal to | `x == 5` |
| `!=` | Not equal | `x != 5` |
| `<` | Less than | `x < 5` |
| `>` | Greater than | `x > 5` |
| `<=` | Less than or equal | `x <= 5` |
| `>=` | Greater than or equal | `x >= 5` |
| `in` | Membership | `"a" in "cat"` |


```python
# If Statement
height = 71   #inches
if height >= 72 :
  print("You are not tall enough for this ride" )

```

```python
# If Else Conditional Statements: two paths
height = 71
if height >= 72 :
  print("You are tall enough for this ride" )
else:
  print("You are not tall enough for this ride")
```

Elif statements allow us to check multiple blocks of code according to which condition is true.

```python
# Elif Statements: multiple categories

birds = 9

if birds < 5:                               # What to do if the first condition is True
  print("I saw less than 5 birds today")
elif birds < 8:                             # another condition of the first is not True
  print("I saw less than 8 birds today")
else:                                       # when other conditions are not True
  print("I saw more than 8 birds today")

```

```python
# Nested conditional statements
age = 11
wants_to_learn_to_swim = True

if wants_to_learn_to_swim == True:
  if age >= 18:
    print("Lets put you in the adult swim class!")
  else:
    print("Lets put you in the children's swim class!")
else:
  print("Please consider taking swimming lessons!")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 3. Loops - Iterating Over Items

the `for` loop allows you to iterate over iterables (like a list) and access each item in it.

The `for` loop is composed of the keyword `for`, followed by the iteration variable, then the word `in`, followed by the iterable (like your list) you want to iterate over, put a colon, then put statements inside the body of the loop, **indented with 4 spaces**.

Let's practice.


```
# Basic Loop DON'T DO THIS (usually)
for i in range(10):
  print(i)

```


```
# DO THIS INSTEAD
cities = ["Oakland", "Berkeley", "San Francisco", "San Jose"]
for city in cities:
    print(city)
```
This isn't just style - it's clearer, less error-prone, and more Pythonic.

"Loop like a native" - Ned Batchelder



The "**break**" statement: stops the entire loop, when the given condition is satisfied.

```python
# code with break
#tell us what position

ages = [17,23,16,34,7,56]
position = 0
for age in ages:
  if age > 7:
    position += 1
  else:
    break

print(f"The person who is {position}th in line is 7 or younger")
```

The "**continue**" statement: stops a specific iteration of a loop, and continues with the next element in the iterable.

```python
# Loop with continue
commute_times = [5,10,15,20,25,30]
count = 0
for i in commute_times:
  if i >= 20:
    continue
  else:
    count += 1
print(f"There are {count} people who have have a commute shorter that 20 minutes")

```

```python
# Loops with strings, conditional statements, and loop control statements (break and continue)!

bike_count = 0
travel_mode = ["bike", "walk", "drive", "walk", "bike"]
for i in travel_mode:
  if i == "bike":
    bike_count += 1
  else:
    continue
print(bike_count)
```

Sometimes, when you are looping, you want both the item and the index of that item. In that case you would want to use `enumerate()`

Use `enumerate()` when you need:



*   To number items (like a ranked list)
*   To know "which one" you're on (first, second, etc.)
* To build output that references position

```python
stations = ["Richmond", "El Cerrito", "Berkeley", "Oakland"]

# THIS IS WRONG: Don't use range(len(...))
print("Using range(len()) - avoid this:")
for i in range(len(stations)):
    print(f"  Stop {i+1}: {stations[i]}")
```

```python
# THIS IS RIGHT: Use enumerate()
print("Using enumerate() - do this:")
for i, station in enumerate(stations):
    print(f"  Stop {i+1}: {station}")
```

```python
# enumerate() with custom start position
print("Numbered list starting at 1:")
for num, station in enumerate(stations, start=1):
    print(f"  {num}. {station}")
```

Another common situation is iterating over multiple lists. For this we want to use `zip()` Fear not, this task is much easier with future libraries we will cover. Nonetheless...

When you have parallel lists (same length, related data), use `zip()`:

```python
cities = ["Oakland", "Berkeley", "SF", "San Jose"]
populations = [433031, 124321, 873965, 1013240]
areas = [78.0, 10.5, 46.9, 180.0]

# THIS IS WRONG: Using indices to access parallel lists
print("Using indices (avoid this):")
for i in range(len(cities)):
    density = populations[i] / areas[i]
    print(f"  {cities[i]}: {density:,.0f}/sq mi")
```

```python
# THIS IS RIGHT: Using zip() to iterate in parallel
print("Using zip() - do this:")
for city, pop, area in zip(cities, populations, areas):
    density = pop / area
    print(f"  {city}: {density:,.0f}/sq mi")
```

### Core Pattern Loops

As you may have deduced from some common examples above, there are certain patterns or types of loops that you will commonly need.

1.   Counting
2.   Acumulating or Summing
3.   Building a new list
4.   Filtering a list
5.   Searching a list



### Pattern 1: Counting

Count items that meet a condition.

```python
# Pattern: COUNTING
# "How many cities have population > 500,000?"

populations = [433031, 124321, 873965, 1013240, 230504]

# Step 1: Initialize counter BEFORE the loop
large_count = 0

# Step 2: Loop and conditionally increment
for pop in populations:
    if pop > 500000:
        large_count += 1  # Same as: large_count = large_count + 1

# Step 3: Use the result AFTER the loop
print(f"Cities with pop > 500K: {large_count}")
```

### Pattern 2: Accumulating (Summing)

Build up a total or running calculation.

```python
# Pattern: ACCUMULATING
# "What's the total population?"

populations = [433031, 124321, 873965, 1013240, 230504]

# Step 1: Initialize accumulator
total = 0

# Step 2: Add each value
for pop in populations:
    total += pop

# Step 3: Use the result
print(f"Total population: {total:,}")
print(f"Average: {total / len(populations):,.0f}")
```

```python
# Conditional accumulation
# "What's the total population of large cities only?"

populations = [433031, 124321, 873965, 1013240, 230504]

large_total = 0

for pop in populations:
    if pop > 500000:
        large_total += pop

print(f"Total population (large cities only): {large_total:,}")
```

### Pattern 3: Building a List

Create a new list/set based on existing data.

```python
# Pattern: BUILDING A LIST
# "Create a list of densities"

cities = ["Oakland", "Berkeley", "SF", "San Jose"]
populations = [433031, 124321, 873965, 1013240]
areas = [78.0, 10.5, 46.9, 180.0]

# Step 1: Start with empty list
densities = []

# Step 2: Append computed values
for pop, area in zip(populations, areas):
    density = pop / area
    densities.append(density)

# Step 3: Use the result
rounded_densities = []
for d in densities:
    rounded_densities.append(round(d))
print(f"Densities: {rounded_densities}")
```

### Pattern 4: Filtering

Select items that meet a condition.

```python
# Pattern: FILTERING
# "Get only the large cities"

cities = ["Oakland", "Berkeley", "SF", "San Jose", "Fremont"]
populations = [433031, 124321, 873965, 1013240, 230504]

# Step 1: Start with empty result list
large_cities = []

# Step 2: Conditionally append
for city, pop in zip(cities, populations):
    if pop > 400000:
        large_cities.append(city)

# Step 3: Use the result
print(f"Large cities (>400K): {large_cities}")
```

### Pattern 5: Finding (Search)

Find the first/best item that meets a condition.

```python
# Pattern: FINDING MAXIMUM
# "Which city has the highest population?"

cities = ["Oakland", "Berkeley", "SF", "San Jose"]
populations = [433031, 124321, 873965, 1013240]

# Step 1: Initialize with first values
max_pop = populations[0]
max_city = cities[0]

# Step 2: Check each and update if better
for city, pop in zip(cities, populations):
    if pop > max_pop:
        max_pop = pop
        max_city = city

# Step 3: Use the result
print(f"Largest city: {max_city} ({max_pop:,})")
```

### Combining Multiple Patterns

Real analysis with complex and nuanced questions will often need to combine multiple patterns in order to extract what.

Let's look at a few examples of this...

```python
# Complex example: Categorize cities by size

cities = ["Oakland", "Berkeley", "SF", "San Jose", "Fremont", "Richmond"]
populations = [433031, 124321, 873965, 1013240, 230504, 116448]

# Build three lists (categorization)
large = []     # > 500K
medium = []    # 200K - 500K
small = []     # < 200K

for city, pop in zip(cities, populations):
    if pop > 500000:
        large.append(city)
    elif pop > 200000:
        medium.append(city)
    else:
        small.append(city)

print(f"Large (>500K):  {large}")
print(f"Medium (200-500K): {medium}")
print(f"Small (<200K):  {small}")
```

```python
# Complex example: Calculate statistics for filtered data

stations = [
    {"name": "Downtown Berkeley", "ridership": 8200, "city": "Berkeley"},
    {"name": "North Berkeley", "ridership": 3100, "city": "Berkeley"},
    {"name": "Ashby", "ridership": 2800, "city": "Berkeley"},
    {"name": "MacArthur", "ridership": 9500, "city": "Oakland"},
    {"name": "19th St", "ridership": 7200, "city": "Oakland"},
]

# Calculate average ridership for Berkeley stations only
berkeley_total = 0
berkeley_count = 0

for station in stations:
    if station["city"] == "Berkeley":
        berkeley_total += station["ridership"]
        berkeley_count += 1

berkeley_avg = berkeley_total / berkeley_count

print(f"Berkeley stations: {berkeley_count}")
print(f"Total ridership: {berkeley_total:,}")
print(f"Average: {berkeley_avg:,.0f}")
```

Often, you may have a large list of files you want to search through to isolate certain files. We can do this using string processing in loops.



```python
# Find files by extension
files = ["data.csv", "report.pdf", "census.csv", "map.png", "transit.csv"]

csv_files = []
for filename in files:
    if filename.endswith(".csv"):
        csv_files.append(filename)

print(f"CSV files: {csv_files}")
```

```python
# Find stations containing a keyword
stations = ["Richmond", "El Cerrito", "North Berkeley", "Downtown Berkeley",
            "Ashby", "MacArthur", "19th St Oakland"]

berkeley_stations = []
for station in stations:
    if "Berkeley" in station:
        berkeley_stations.append(station)

print(f"Berkeley stations: {berkeley_stations}")
```

We've covered a lot so far regarding loops, but here are a few common mistakes to avoid when writing loops:



1.   **Mistake 1**: Initializing inside the loop
2.   **Mistake 2**: Using wrong variable (index vs item)
3.   **Mistake 3**: Forgetting to use the result after the loop


```python
# MISTAKE 1: Initializing inside the loop
numbers = [1, 2, 3, 4, 5]

# WRONG - resets every iteration!
for num in numbers:
    total = 0  # BUG: Resets to 0 each time!
    total += num

print(f"Wrong total: {total}")  # Only gets last number!

# RIGHT - initialize before the loop
total = 0
for num in numbers:
    total += num

print(f"Correct total: {total}")
```

```python
# MISTAKE 2: Using wrong variable (index vs item)
cities = ["Oakland", "Berkeley", "SF"]

# CONFUSING - mixing patterns
# for i in cities:  # 'i' suggests index, but it's actually the city!
#     print(i)

# CLEAR - name matches what it is
for city in cities:
    print(city)
```

```python
# MISTAKE 3: Forgetting to use the result after the loop
numbers = [10, 20, 30]

total = 0
for num in numbers:
    total += num
    print(f"Total so far: {total}")  # Printing inside loop

# The final result should be used AFTER the loop
print(f"Final total: {total}")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Summary: When to Use What

| Situation | Use This |
|-----------|----------|
| Process each item | `for item in items:` |
| Need position/index | `for i, item in enumerate(items):` |
| Parallel lists | `for a, b in zip(list_a, list_b):` |
| Dictionary key-values | `for key, value in dict.items():` |
| Count occurrences | Initialize counter, `+= 1` in loop |
| Sum values | Initialize total, `+=` in loop |
| Build new list | Initialize `[]`, `.append()` in loop |
| Build new set | Initialize `set()`, `.add()` in loop |
| Filter items | Initialize `[]`, conditionally `.append()` |

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Loop Tracing

Being able to trace through a loop is a crucial skill. Let's practice.



```python
# Trace this loop step by step
values = [10, 25, 15]
total = 0

for val in values:
    total = total + val
    print(f"  After adding {val}: total = {total}")

print(f"Final total: {total}")
```

---

### Trace Table Exercise

Before running the next cell, fill in this trace table:

```
numbers = [2, 5, 3]
result = 1

| Iteration | num | result (before) | result (after) |
|-----------|-----|-----------------|----------------|
| 1         | 2   | 1               | ____            |
| 2         | 5   | ____             | ____            |
| 3         | 3   | ____             | ____            |
```

Record your trace **without running the loop**:

- `trace_before`: list of `result` values *before* each of the 3 iterations
- `trace_after`: list of `result` values *after* each iteration
- `trace_final`: the final value of `result`

```python
trace_before = ...
trace_after = ...
trace_final = ...
```

```python
grader.check("trace")
```

```python
# Check your trace
numbers = [2, 5, 3]
result = 1

for num in numbers:
    print(f"Before: result = {result}, multiplying by {num}")
    result = result * num
    print(f"After: result = {result}")
    print()

print(f"Final: {result}")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 4. Boolean Expressions

Jumping back to our conditionals we worked on before loops, how can we make our slightly verbose if-elif-else's better? Boolean Expressions using boolean operators.

Boolean Operators:

*  **Or** = extends our condition to include atleast one of the given conditions
*  **And** = Narrows the search by requiring all of the given conditions
*  **Not** = exclude from the search according to a specific condition

Critically to understanding: `not`inverts the true value of a boolean expression.



```python
# 'Or' - atleast one condition must be valid

building_height_stories = 10
distance_from_property_line = 15
max_height = 8
min_setback = 20

if building_height_stories > max_height or distance_from_property_line < min_setback:
    print("Project requires a zoning variance.")
else:
    print("Project compliant with zoning.")
```

```python
# "And" - evaluates only if both conditions are true

dist_to_transit = 0.4
is_near_commercial_zone = True


if dist_to_transit < 0.5 and is_near_commercial_zone == True:
  print("TOD is suitable for this site")
else:
  print("TOD is not suitable for this site")

```

```python
# 'Not' - exlude from search according to the condition

speed_limit = 40
has_bike_lane = False
traffic_volume = 15000

if speed_limit >= 35 and not has_bike_lane and traffic_volume > 10000:
    print("Recommend: High-priority bike lane installation.")
else:
    print("No immediate bike lane action needed.")
```

CRITICALLY, the code above runs because `has_bike_lane` is True, because the `not` operator inverted its value.

```python
# Combining conditions with and, or, not
population = 150000
has_rail = True

# Both conditions must be true
if population > 100000 and has_rail:
    print("Medium city with rail transit")

# At least one condition must be true
if population > 500000 or has_rail:
    print("Major city OR has rail")

# Negation
if not has_rail:
    print("No rail service")
```

Using the previous examples, we can also combine several boolean operators to evaluate multiple conditions:

```python
# Combining boolean operators

dist_to_transit = 0.4
is_near_commercial_zone = True
parcel_size = 0.6

if (dist_to_transit < 0.5 and parcel_size <= 0.5) or is_near_commercial_zone == True:
  print("TOD is suitable for this site")
else:
  print("TOD is not suitable for this site")

```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Exercises - Lists, Conditional, Loops, and Boolean Expressions

At this stage of your python learning, it is essential that you can look at code using these basic operators and keywords and be able to explain in your own words what is happening behinds the scenes.

For each of the exercises below, explain in your own words the conditional and looping logic of your code in a text block below your code, similar to how we traced the logic of our loops.

---

### Lists

Create an empty list as a variable `my_list`. Practice populating and depopulating the list by using each of the list methods we learned earlier.  

```python
# Code here
my_list = ...
```

```python
grader.check("q1")
```

---

### Conditionals

Using an if-else statement, calculate the cost of a bike lane installation for a street using these rules, choosing the type of bike lane you would like installed on the street:

**Fees:**
- Base administrative fee: $ 100000

* Class III Lane: + $ 41000

* Class II Lane: +  98000

* Class IV Lane: + $ 280000

* Class I Lane: + $ 430000

- Cost per square foot of asphalt: $ 17

**Your code must:**
1. Set exactly one of the four boolean lane variables to `True` (the others `False`) to choose which lane type to build -- your code will be checked to confirm only one lane type is selected
2. Use the given variables (don't modify them)
3. Calculate total cost
4. Print formatted breakdown

Store the final result in a variable named `total_cost`.

```python
# Given variable - DO NOT MODIFY
project_area_sqft = 34000

# Boolean Lane variable - MODIFY ACCORDING TO WHAT YOU WANT TO BUILD
# Set exactly ONE of these to True
Class_III_lane = ...
Class_II_lane = ...
Class_IV_lane = ...
Class_I_lane = ...

# Other variables
base_fee = 100000
asphalt_cost = 17 * project_area_sqft

# Write your code below
if Class_III_lane:
    lane_fee = ...
elif Class_II_lane:
    lane_fee = ...
elif Class_IV_lane:
    lane_fee = ...
else:
    lane_fee = ...

total_cost = ...
print(f"Total: ${total_cost}")
```

```python
grader.check("q2")
```

If your project budget is $ 1000000, write an if-else statement that checks if your project is overbudget. For each of your conditions write a print statement that prints whether the project is overbudget or not.

Store the comparison result in a boolean variable named `over_budget`.

```python
budget = 1000000
over_budget = ...
if over_budget:
    print(...)
else:
    print(...)
```

```python
grader.check("q3")
```

---

### Loops

we have arrived at an iconic moment in your Python learning journey! **The FizzBuzz Problem!** This is a classic exercise that lets us practice looping, operators, and conditionals.

Write a `for` loop that iterates through the numbers 1 to 100 using `range()` and prints each number. However, the program must follow these specific conditions:

*  If the number is a multiple of both 3 and 5, print "FizzBuzz" instead of the number.
*  If the number is a multiple of 3 only, print "Fizz".
* If the number is a multiple of 5 only, print "Buzz".
* Otherwise (if it is a multiple of neither 3 nor 5), print the number itself.

As you loop, also append each printed value to a list named `fizzbuzz_results` (append numbers as numbers and the words as strings).

```python
# Your code here
fizzbuzz_results = ...
```

```python
grader.check("q4")
```

<!-- BEGIN QUESTION -->

---

### Write your own problem using loops and conditionals!

Much like our rubber duck method that we learned about, a great way to really make sure you understand a programming concept is to try to write your own problem. Give it a try!

```python
# Your code here
```

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload your submission zip to **LAB04 - Python Basics II** on [Gradescope](https://www.gradescope.com/).

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1 (points: 1)
```
>>> assert isinstance(my_list, list)
```


Tests for q2 (points: 2)
```
>>> assert sum([Class_III_lane, Class_II_lane, Class_IV_lane, Class_I_lane]) == 1
```
```
>>> assert total_cost in (719000, 776000, 958000, 1108000)
```


Tests for q3 (points: 1)
```
>>> assert isinstance(over_budget, bool)
```
```
>>> assert over_budget == (total_cost > 1000000)
```


Tests for q4 (points: 2)
```
>>> assert len(fizzbuzz_results) == 100
```
```
>>> assert fizzbuzz_results[0] == 1 and fizzbuzz_results[2] == 'Fizz'
```
```
>>> assert fizzbuzz_results[4] == 'Buzz' and fizzbuzz_results[14] == 'FizzBuzz'
```


Tests for trace (points: 1)
```
>>> assert trace_before == [1, 2, 10]
```
```
>>> assert trace_after == [2, 10, 30]
```
```
>>> assert trace_final == 30
```
